Cellular characterization of the mouse collecting lymphatic vessels reveals that lymphatic muscle cells are the innate pacemaker cells

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Recommendations for the authors):

The authors have done an impressive job in responding to the previous critique and even gone beyond what was asked. I have only very minor comments on this excellent manuscript. The manuscript also needs some light editing for grammar and readability.

We have worked to improve the grammar and readability of the manuscript.

Comments:

Lines 227-234: At what age was tamoxifen administered to the various CreERTM mice?

We have updated the ages of the mice used in this study in the methods sections.

UMAP in Figure 5A is missing label for cluster 19.

The UMAP in Figure 5A has the label for cluster 19 at the center-bottom of the image.

Supplement Figure 6: Cluster 10 seems to be separate from the other AdvC clusters, and it includes some expression of Myh11 and Notch3. Further, there is low expression of Pdgfra in this cluster, which can be seen in panel B and panels D-I. Are the Pdgfra negative cells in the pie charts from cluster 10? Could the cells in this cluster by more LMC like than AdvC like?

We agree with the reviewer that the subcluster 10 of the fibroblasts cells are intriguing if only a minor population. When assessing just this population of cells, which is 77 cells out of 2261 total, 40 of the 77 were Pdgfra+ and of the 37 remaining Pdgfra- but 11 of those were still CD34+. Thus at least half of these cells could be expected to have the PdgfraCreERTM. Only 8 of the 37 were Pdgfra-Notch3+ while 12 cells were Pdgfra+Notch3+, and only 3 were Pdgfra-Myh11+ while 3 were Pdgfra+Myh11+. 26 of 77 cells were Pdgfra+Pdgfrb+ double positive, while 12 of 37 Pdgfra- cells were still Pdgfrb+. Additionally, within the 77 cells of subcluster 10 17 were positive for Scn3a (Nav1.3), 21were positive for Kcnj8 (Kir6.1), and 33 were positive for Cacna1c (Cacna1c) which are typically LMC markers would support the reviewers thinking that this group contains a fibroblast-LMC transitional cell type. Only 2 of 77 cells were positive for the BK subunit (Kcnma1), which is a classic smooth muscle marker. Another possibility is this population represents the Pdgfra+Pdgfrb+ valve interstitial cells we identified in our IF staining and in our reporter mice. Of note almost all cells in this cluster were Col3a1+ and Vim+. Even though we performed QC analysis to remove doublets, it is also possible some of these cells could represent doublets or contaminants, however the low % of Myh11 expression, a very highly expressed gene in LMCs especially compared to ion channels, would suggest this is less likely. Assessing the presence of this particular cell cluster in future RNAseq or with spatial transcriptomics will be enlightening.

Line 360. Proofread section title.

We have simplified this title to read “Optogenetic Stimulation of iCre-driven Channel Rhodopsin 2”

Lines 370-371. Are the length units supposed to be microns or millimeters?

We have corrected this to microns as was intended. Thank you for catching this error.

The resolution for each UMAP analysis should be stated, particularly for the identification of subclusters. How was the resolution chosen?

To select the optimal cluster resolution, we used Clustree with various resolutions. We examined the resulting tree to identify a resolution where the clusters were well-separated and biologically meaningful, ensuring minimal merging or splitting at higher resolutions. Our goal was to find a resolution that captures relevant cell subpopulations while maintaining distinct clusters without excessive fragmentation. We have now stated the resolution for the subclustering of the LECs, LMCs, and fibroblasts. We have also added greater detail regarding the total number of cells, QC analysis, and the marker identification criteria used to the methods sections. We used resolution of 0.5 for sub-clustering LMCs, 0.87 for LECs, and 1.0 for fibroblasts. These details are now added to the manuscript.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Collecting lymphatic vessels (cLVs) exhibit spontaneous contractions with a pressure-dependent frequency, but the identity of the lymphatic pacemaker cell is still debated. By analogy to pacemakers in the GI and lower urinary tracts, proposed cLV pacemaker cells include interstitial cells of Cajal like cells (ICLC) or the lymphatic muscle (LMCs) cells themselves. Here we combined immunofluorescence and scRNAseq analyses with electrophysiological methods to examine the cellular constituents of the mouse cLV wall and assess whether any cell type exhibited morphological and functional processes characteristic of pacemaker cells: a continuous if not contiguous network integrated into the electrical syncytium; spontaneous Ca2+ transients; and depolarization-induced propagated contractions. We employed inducible Cre (iCre) mouse models routinely used to target these specific cell populations including: c-kitCreERT2 to target ICLC; PdgfrβCreERT2 to target pericyte-like cells; PdgfrαCreERTM to target CD34+ adventitial cells and ICLC; and Myh11CreERT2 to target LMCs directly. These specific inducible Cre lines were crossed to the fluorescent reporter ROSA26mT/mG, the genetically encoded Ca2+ sensor GCaMP6f, and the light-activated cation channel rhodopsin2 (ChR2). c-KitCreERT2 labeled both a sparse population of LECs and round adventitial cells that responded to the mast cell activator compound 48-80. PdgfrβCreERT2 drove recombination in both adventitial cells and LMCs, limiting its power to discriminate a pericyte-specific population. PdgfrαCreERTM labeled a large population of interconnected, oak leaf-shaped cells primarily along the adventitial surface of the vessel. Of these cells, only LMCs consistently, but heterogeneously, displayed spontaneous Ca2+ events during the diastolic period of the contraction cycle, and whose frequency was modulated in a pressure-dependent manner. Optogenetic depolarization through the expression of ChR2 under control of Myh11CreERT2, but not PdgfrαCreERTM or c-KitCreERT2, resulted in propagated contractions upon photo-stimulation. Membrane potential recordings in LMCs demonstrated that the rate of diastolic depolarization significantly correlated with contraction frequency. These findings support the conclusion that LMCs, or a subset of LMCs, are responsible for mouse cLV pacemaking.

Revised version: Collecting lymphatic vessels (cLVs) exhibit spontaneous contractions with a pressure-dependent frequency, but the identity of the lymphatic pacemaker cell is still debated. Here we combined immunofluorescence and scRNAseq analyses with electrophysiological methods to examine the cellular constituents of the mouse cLV wall and assess whether any cell type exhibited morphological and functional processes characteristic of pacemaker cells. We employed inducible Cre (iCre) mouse models to target specific cell populations including c-kitCreERT2 to target interstitial cells of Cajal like cells; PdgfrβCreERT2 to target pericyte-like cells; PdgfrαCreERTMto target CD34+ adventitial cells; and Myh11CreERT2to target lymphatic muscle cells (LMCs) directly. These inducible Cre lines were crossed to the fluorescent reporter ROSA26mT/mG, the genetically encoded Ca2+ sensor GCaMP6f, and the light-activated cation channel rhodopsin2 (ChR2). Only LMCs consistently, but heterogeneously, displayed spontaneous Ca2+ events during the diastolic period of the contraction cycle, and whose frequency was modulated in a pressure-dependent manner. Further, optogenetic depolarization with ChR2 only induced propagated contractions in LMCs. Membrane potential recordings in LMCs demonstrated that the rate of diastolic depolarization significantly correlated with contraction frequency. These findings support the conclusion that LMCs, or a subset of LMCs, are responsible for mouse cLV pacemaking.

Introduction

The spontaneous contractions of collecting lymphatic vessels (cLV) are an integral component to fluid and macromolecule homeostasis as they provide the force to transport fluid from the interstitial spaces back to the blood circulation (Scallan et al., 2016). In humans, spontaneous contractile activity is estimated to account for 2/3 of lymph transport (Engeset et al., 1977) and this function is significantly compromised in patients suffering from lymphedema, whose cLVs typically display weak and irregular or entirely absent contractile activity (Olszewski, 2002). Ex vivo studies, in which the intraluminal pressure can be precisely controlled, have refined our understanding of the pressure-dependent regulation of contraction frequency (Benoit et al., 1989; Gashev et al., 2004), with some mouse cLVs displaying a 10-fold increase in contraction frequency over a 10 cmH2O pressure gradient (Scallan and Davis, 2013; Zawieja et al., 2018a). The observation that cLVs, often cannulated at various lengths for ex vivo preparations, retain a consistently tunable contraction frequency points to the presence of (a) pacemaker cell(s) innate to the structure of the cLV wall and with a seemingly ubiquitous presence along the length of the vessel (Zawieja et al., 1993; Castorena-Gonzalez et al., 2018b). Furthermore, isolated cLVs typically display single pacemaker initiation sites unless damaged or electrically uncoupled by pharmacological inhibition of gap junctions or genetic deletion of Gjc1 (Connexin 45, Cx45) (Behringer et al., 2017; Castorena-Gonzalez et al., 2018b; Castorena-Gonzalez et al., 2020). In sum, this suggests the pacemaker cell(s) is(are) likely both ubiquitous and continuous, to allow for electrical conduction via gap junctions, along the length of the cLV and prevent colliding contractile waves which would impair lymph transport.

Investigations into the cLV pacemaker identity have focused largely on cells termed interstitial cells of Cajal like cells (ICLC; or telocytes) (McCloskey et al., 2002; Briggs Boedtkjer et al., 2013), as they display some morphological and cell marker expression profiles similar to the interstitial cells of Cajal (ICC), which are bona fide pacemakers in the gastrointestinal (GI) tract. ICC are classically identified by either methylene blue staining and expression of CKIT, and coordinate GI smooth muscle contraction (Maeda et al., 1992; Ward et al., 1994; Ordog et al., 1999). ICC also express the canonical Ca2+ activated chloride channel Anoctamin 1 (Ano1) (Gomez-Pinilla et al., 2009), which is required for pacemaker activity (Hwang et al., 2009; Zhu et al., 2009; Singh et al., 2014). Previous reports in sheep mesenteric lymphatic vessels identified a population of cKIT+, VIMENTIN+, ICLC in the vessel wall between the endothelial and LMC layer and running along the axis of the vessel (McCloskey et al., 2002). Investigations in the human thoracic duct also identified a significant population of ICLC in close proximity to the lymphatic muscle cells (LMCs) evident by methylene blue staining, immunostaining for CD34, VIMENTIN, and cKIT, as well as the gold standard of electron microscopy (Briggs Boedtkjer et al., 2013). However, neither study could determine if these cells had functional electrical communication with the LMCs or demonstrate either a membrane electrical clock or internal Ca2+ clock to drive the rhythmic lymphatic vessel contractions observed ex vivo. LMCs share a functional similarity to ICC in that they also display the Ano1 mediated Ca2+ activated chloride current (Van Helden, 1993; Toland et al., 2000; Mohanakumar et al., 2018) (Zawieja et al., 2019), that regulates pacemaking. Spontaneous transient depolarizations, presumably Ano1 dependent, were recorded in mesenteric cLVs from guinea pigs (Van Helden, 1993; von der Weid et al., 2008) providing a mechanism for membrane potential instability to drive AP initiation. Furthermore, computational models have proposed LMC sarcoplasmic reticulum (SR) Ca2+ release as the oscillator mechanism driving pacemaking (Imtiaz et al., 2007). SR Ca2+ release has also been implicated in pericyte regulation of arterioles (Hashitani et al., 2015; van Helden and Imtiaz, 2019), in microvascular vasomotion (Boedtkjer et al., 2008; Aalkjaer et al., 2011; van Helden and Imtiaz, 2019), and in the contraction waves of atypical muscle cells of the lower urinary tract (Grainger et al., 2022).

Earlier version: Presently, no investigations have clearly identified the cellular identities of possible pacemaker cells within the cLVs of the mouse. Mouse cLVs exhibit contractile parameters and conduction speed equivalent to those of human vessels (Castorena-Gonzalez et al., 2018b) and their simplified architecture, compared to larger mammals, in combination with the genetic tools developed for the mouse model, allowed us to test for a fundamental pacemaker cell in the cLV. In this study we utilized multiple genetic mouse models, confocal imaging of fluorescent reporters, cell specific expression of GCaMP6f for Ca2+ imaging, and optogenetic light-activated depolarization to both visualize and test the functional aspects of putative pacemaker cells, and membrane potential recordings in LMCs in pressure-challenged cLVs. We also performed immunostaining and single cell RNA sequencing (scRNAseq) of isolated cLVs to provide greater detail to the heterogenous cellular populations found within the mouse cLVs. Despite identifying a significant population of CD34+Pdgfrα+ adventitial cells along the length of mouse cLVs, the results of our functional studies support a myogenic origin of pacemaking in cLVs.

Revised version: Presently, no investigations have clearly identified the cellular identities of possible pacemaker cells within the cLVs of the mouse. Mouse cLVs exhibit contractile parameters and conduction speed equivalent to those of human vessels (Castorena-Gonzalez et al., 2018b) and their simplified architecture, compared to larger mammals, in combination with the genetic tools developed for the mouse model, allowed us to test for a fundamental pacemaker cell in the cLV. In this study we utilized multiple genetic mouse models, confocal imaging of fluorescent reporters, cell specific expression of GCaMP6f for Ca2+ imaging, and optogenetic light-activated depolarization to both visualize and test the functional aspects of putative pacemaker cells, along with membrane potential recordings in LMCs in pressure-challenged cLVs. We also performed immunostaining and single cell RNA sequencing (scRNAseq) of isolated cLVs to provide greater detail to the heterogenous cellular populations found within the mouse cLVs. Despite identifying a significant population of CD34+Pdgfrα+ adventitial cells along the length of mouse cLVs, the results of our functional studies support a myogenic (LMC) origin of pacemaking in cLVs.

Results

Methylene Blue Staining Reveals a Minor Population of Cells in Mouse cLVs

Earlier version: Methylene blue staining was used to identify an ICLC population in the human lymphatic thoracic duct (Briggs Boedtkjer et al., 2013). In our isolated and cleaned lymphatic inguinal axillary collecting vessels (IALVs), methylene blue stained a significant number of cells with variable density along the length of the IALV and heterogenous cell morphologies (Figure 1A-C). A significant portion of the stained cells resembled lymphatic vessel-associated macrophages with an elongated shape, while other cells were smaller and circular (Figure 1D-F). Methylene blue also appeared to stain mast cells as there were large ovoid cells on the adventitia of the vessel with intracellular granules. In addition, methylene blue stained a minor population of cells that exhibited long and thin axon-like extensions which appeared to have a slight helical orientation, with a small central body and nucleus (Figure 1C). None of these cell populations were aligned with the longitudinal axis of the vessel that would permit efficient coupling or regulation across the circumferential LMCs required for coordinated propagation along the length of the vessel.

Revised version: Methylene blue staining was used to identify an ICLC population in the human lymphatic thoracic duct (Briggs Boedtkjer et al., 2013). In our isolated and cleaned lymphatic inguinal axillary collecting vessels (IALVs), methylene blue stained a significant number of cells with variable density along the length of the IALV and heterogenous cell morphologies (Figure 1A-C). A significant portion of the stained cells resembled lymphatic vessel-associated macrophages with an elongated shape, while other cells were smaller and circular (Figure 1D-F). Methylene blue also appeared to stain mast cells as there were large ovoid cells with intracellular granules on the adventitial surface of the vessel. In addition, methylene blue stained a minor population of cells that exhibited long and thin axon-like extensions which appeared to have a slight helical orientation, with a small central body and nucleus (Figure 1C). None of these cell populations were aligned with the longitudinal axis of the vessel that would permit efficient coupling or regulation across the circumferential layer of LMCs required for coordinated propagation along the length of the vessel.

Methylene blue staining of isolated mouse IALVs

Representative image of an isolated and cleaned IALV after methylene blue staining which revealed cells of various morphology. (B) is the zoomed in image of the yellow dotted box in A which contained large ovoid cells with granular staining (B, yellow asterisks). Fine cellular extensions (red asterisks) stained by methylene blue in some cells were visualized with color channel separation and division (C). (D, E) Similar as B and C, but in a separate vessel which stained with a higher density of methylene blue stained cells some of which had limited cellular processes. F) Focal reconstruction from imaging a methylene blue stained IALV using an upright microscope and immersion objective.

Immunofluorescence Imaging of IALVs Stained for ICLC, LEC, and LMC Markers

Earlier version: We next stained IALVs for the putative telocyte/ICLC markers cKIT, CD34, and the intermediate filament Vimentin, which have been previously utilized for ICLC identification in human and sheep lymphatic tissues (McCloskey et al., 2002; Briggs Boedtkjer et al., 2013). Additionally, an antibody to the intermediate filament Desmin was used to label muscle cells (McCloskey et al., 2002). IALVs stained with cKIT (Figure 2B) showed robust signal in large ovoid cells with a non-segmented circular nucleus (Figure 2A), characteristic of mast cells that were located in the outer part of the adventitia. Similarly, cKIT stained populations of elongated cells as well as circular cells with variable densities throughout the IALV wall, similar to methylene blue+ cell populations (Figure 2B, J). Staining for CD34 revealed a large population of cells that were seemingly contiguous along the length of the vessel. The CD34+ cells generally had multiple lobular processes and a “oak leaf” like appearance, typically a characteristic of fibroblasts, though some contained short, thin dendrite-like extensions (Figure 2C, G, K). The CD34+ cells were negative for Desmin [Figure 2H), which primarily stained the circumferential LMCs (Figure 2F; note that the largely non-circumferential cell organization in this region is typical for a lymphatic endothelial valve site (Bridenbaugh et al., 2013a)]. Furthermore, CD34+ cells and cKIT+ cells were separate populations (Figure 2D, L). A Vimentin antibody labeled lymphatic endothelial cells (LECs) which exhibited a horizontal cobblestone morphology in parallel with the vessel axis (Figure 2E, I), while also co-labeling the majority of the CD34+ cells (Figure 2H) and cKIT+ cells (Figure 2L). Videos of the half vessel z-stacks are provided (Supplemental Movies 1-3 for Figure 2D, H, and L respectively).

Revised version: We next stained IALVs for the putative telocyte/ICLC markers cKIT, CD34, and the intermediate filament VIMENTIN, which have been previously utilized for ICLC identification in human and sheep lymphatic tissues (McCloskey et al., 2002; Briggs Boedtkjer et al., 2013). Additionally, an antibody to the intermediate filament Desmin was used to label muscle cells (McCloskey et al., 2002). IALVs stained for cKIT (Figure 2B) showed robust signal in large ovoid cells with a non-segmented circular nucleus (Figure 2A), characteristic of mast cells that were in the outer part of the adventitia. Similarly, cKIT stained populations of elongated cells as well as circular cells with variable densities throughout the IALV wall, similar to methylene blue+ cell populations (Figure 2B, J). Staining for CD34 revealed a large population of cells that were seemingly contiguous along the length of the vessel. The CD34+ cells generally had multiple lobular processes and a “oak leaf” like appearance, typically a characteristic of fibroblasts, though some contained short, thin dendrite-like extensions (Figure 2C, G, K). The CD34+ cells were negative for Desmin [Figure 2H), which primarily stained the circumferential LMCs (Figure 2F; note that the largely non-circumferential cell organization in this region is typical for a lymphatic endothelial valve site (Bridenbaugh et al., 2013a)]. Furthermore, CD34+ cells and cKIT+ cells were separate populations (Figure 2D, L). A VIMENTIN antibody labeled lymphatic endothelial cells (LECs) which exhibited a horizontal cobblestone morphology in parallel with the vessel axis (Figure 2E, I), while also co-labeling the majority of the CD34+ cells (Figure 2H) and cKIT+ cells (Figure 2L). Videos of the half vessel z-stacks are provided (Supplemental Movies 1-3 for Figure 2D, H, and L respectively).

Staining Mouse IALVs for ICLC Markers

Representative immunofluorescent max projections of half vessel confocal image stacks imaged from mouse IALVs stained for ICLC markers. DAPI (A), cKIT (B), and CD34 (C) and their merged image (D). Representative max projections of the intermediate filament VIMENTIN (E), the intermediate filament desmin (F), CD34 (G) and their merged image (H). Representative max projection of VIMENTIN (I), cKIT (J), CD34 (K) and their merged image (L). Scale bar = 100 µm for all images.

Earlier version: Of the cells stained in Figure 2, the CD34+ population was intriguing due to its high density and distribution throughout the length of the IALV, which potentially would be conducive to effective regulation of LMC excitability. In addition to CD34 and Vimentin, PDGFRα staining is also commonly ascribed to both telocytes (Vannucchi et al., 2013; Xiao et al., 2013; Zhou et al., 2015) as well as fibroblasts (Kimura et al., 2021; Clayton et al., 2022). We performed immunofluorescence imaging for PGDFRα counterstained with CD34 and markers for LMCs, LECs, and pericytes. As noted in Figure 2, CD34+ cells (Figure 3A) did not co-label LMCs (Figure 3D) which were smooth muscle actin+ (SMA, Figure 3B) and Calponin+ (Figure 3C). However, nearly all CD34+ (Figure 3E) cells were also PDGFRα+ (Figure 3F, H). Occasionally some overlap of PDGFRα and SMA+ signal was noted (Figure 3G, H). LECs staining with CD31 (PECAM, Figure 3I) revealed the expected rectangular elongated cobblestone morphology that was distinct from the PDGFRα+ cells (Figure 3J, L). Staining for Calponin also specifically labeled LMCs (Figure 3K) but not PDGFRα+ cells (Figure 3L). Lastly, we stained for PDGFRα, CD34, and PDGFRβ, a commonly used pericyte marker (Figure 3 M-P). As above, CD34 and PDGFRα were highly colocalized (Figure 3Q, R, T), and many of the CD34+ and PDGFRα+ cells were also PDGFRβ+ (Figure 3P). PDGFRβ also stained some circumferential LMCs (Figure 3Q). During the imaging of mouse IALVs for these markers, we also observed that the lymphatic secondary endothelial valves were populated by elongated cells that stretched the length of the valve leaflet and were positive for CD34, PDGFRα, and PDGFRβ, with varying intensities. These cells could be observed in most, if not all, the valves we assessed and found within both leaflets of the valve (Figure 3R,S). These cells had long, thin extensions that were branched, along with apparent dendrite extensions with a morphology that closely resembled those described of pericytes or telocytes (Popescu and Faussone-Pellegrini, 2010). PDGFRα+ or CD34+ cells with this morphology were only observed in the valve leaflets, and thus seemed insufficient to regulate pacemaking as normal contractions are observed in cLVs without secondary valves (Van Helden, 1993; Gashev et al., 2002). Representative z-stacks demonstrating these valve-located “telocyte” shaped cells (Figure 3R,S) are provided as Supplemental Movies 4 and 5.

Revised version: Of the cells stained in Figure 2, the CD34+ population was intriguing due to its high density and distribution throughout the length of the IALV, which potentially would be conducive to effective regulation of LMC excitability. In addition to CD34 and VIMENTIN, PDGFRα staining is also commonly ascribed to both telocytes (Vannucchi et al., 2013; Xiao et al., 2013; Zhou et al., 2015) as well as fibroblasts (Kimura et al., 2021; Clayton et al., 2022). We performed immunofluorescence imaging for PGDFRα counterstained with CD34 and markers for LMCs, LECs, and pericytes. As noted in Figure 2, CD34+ cells (Figure 3A) did not co-label LMCs (Figure 3D) which were smooth muscle actin+ (SMA, Figure 3B) and Calponin+ (Figure 3C). However, nearly all CD34+ (Figure 3E) cells were also PDGFRα+ (Figure 3F, H). Occasionally some overlap of PDGFRα and SMA+ signal was noted (Figure 3G, H). LECs staining with CD31 (PECAM, Figure 3I) revealed the expected rectangular elongated cobblestone morphology that was distinct from the PDGFRα+ cells (Figure 3J, L). Staining for Calponin also specifically labeled LMCs (Figure 3K) but not PDGFRα+ cells (Figure 3L). Lastly, we stained for PDGFRα, CD34, and the commonly used pericyte marker PDGFRβ (Figure 3 M-P). As above, CD34 and PDGFRα were highly colocalized (Figure 3Q, R, T), and many of the CD34+ and PDGFRα+ cells were also PDGFRβ+ (Figure 3P). PDGFRβ also stained some circumferential LMCs (Figure 3Q). During the imaging of mouse IALVs for these markers, we also observed that the lymphatic secondary endothelial valves were populated by elongated cells that stretched the length of the valve leaflet and were positive for CD34, PDGFRα, and PDGFRβ, with varying intensities. These cells could be observed in most, if not all, the valves we assessed and found within both leaflets of the valve (Figure 3R,S). These cells had long, thin extensions that were branched, along with apparent dendrite-like extensions with a morphology that closely resembled those described for pericytes or telocytes (Popescu and Faussone-Pellegrini, 2010). PDGFRα+ or CD34+ cells with this morphology were only observed in the valve leaflets, and thus seemed insufficient to regulate pacemaking as normal contractions are observed in cLVs without secondary valves (Van Helden, 1993; Gashev et al., 2002). Representative z-stacks demonstrating these valve-located “telocyte” shaped cells (Figure 3R,S) are provided as Supplemental Movies 4 and 5.

Immunofluorescence Labeling of Mouse IALVs with Markers for ICLC, LMC, LEC, and Immune Cell Populations

We stained isolated mouse IALVs with cellular markers used to differentiate various cell types observed in cLVs. Half vessel image stacks were taken with confocal microscopy and the resulting representative max projections are shown. (A) CD34 stained cells and LMC staining with SMA (B) and calponin (C) and the corresponding merged (D) image. There was significant overlap in (E) CD34 staining along with the fibroblast marker PDGFRα compared to LMC staining with SMA (G) and the merged (H) image. The endothelial marker CD31 (I) to delineate LECs with PDGFRα staining (J), and the LMC marker calponin (K) with the merged image (L) revealed 3 separate populations of cells. PDGFRβ (O) stained many cells that were CD34 (M) and PDGFRα (N) positive, as seen in the merge imaged (P), in addition to PDGFRβ signal detected in the LMC layer (Q). Max projections of only the luminal frames of a z-stack at lymphatic valve locations revealed PDGFRβ, CD34, and PDGFRα labeling in bipolar shaped cells with long extensions that traveled throughout the valve leaflets (V, W). d Control IALV (Y) stained only with secondary antibody. Scale bar = 100 µm for all images.

Earlier version: We next determined the degree of colocalization between the CD34 and PDGFRα signal given the significant overlap in their staining profile. Colocalization analysis of PDGFRα (SuppFigure 1A) and CD34 (SuppFigure 1B) and their colocalization (SuppFigure 1C) was determined with the FIJI BIOP-JACoP tool. The Pearson’s coefficient was 0.83 (SuppFigure 1D) and Mander’s coefficient of overlap 0.80 was for the PDGFRα+ signal and 0.87 for the CD34 signal (SuppFigure 1E). Colocalization between Myh11 and PDGFRα was significantly lower (SuppFigure 1D-F) with a Pearson’s coefficient of 0.30 (SuppFigure 1G), whereas the Mander’s coefficient for Myh11 overlap with PDGFRα was 0.077 and PDGFRα signal overlap with Myh11 was 0.043 (SuppFigure 1H). The high degree of colocalization CD34 and PDGFRα signal informed our use of the commercially available transgenic PdgfrαCreERTM mouse model to target these cells. The vast majority of the PDGFRα+ cells were located in the adventitial layer (SuppFigure 2A-D), which varied between 1-3 PDGFRα+ cells thick (SuppFigure 2E). Under this layer, we observed only a single layer of largely circumferential LMCs stained by Myh11 (SuppFigure 2B) sitting atop a single layer of CD31+ LECs (SuppFigure 2A). We also observed occasional PDGFRα+ cells or their extensions located in the sub-endothelial space (SuppFigure 2 E’, E”) positioned between the LECs and the LMCs.

Revised version: We next determined the degree of colocalization between the CD34 and PDGFRα signal given the significant overlap in their staining profile. Colocalization analysis of PDGFRα (SuppFigure 1A) and CD34 (SuppFigure 1B) and their colocalization (SuppFigure 1C) was determined with the FIJI BIOP-JACoP tool. The Pearson’s coefficient was 0.83 (SuppFigure 1D) and Mander’s coefficient of overlap was 0.80 for the PDGFRα+ signal and 0.87 for the CD34 signal (SuppFigure 1E). Colocalization between Myh11 and PDGFRα was significantly lower (SuppFigure 1D-F) with a Pearson’s coefficient of 0.30 (SuppFigure 1G), whereas the Mander’s coefficient for Myh11 overlap with PDGFRα was 0.077 and 0.043 for PDGFRα signal overlap with Myh11 (SuppFigure 1H). The high degree of colocalization CD34 and PDGFRα signal informed our use of the commercially available transgenic PdgfrαCreERTM mouse model to target these cells. The vast majority of the PDGFRα+ cells were located in the adventitial layer (SuppFigure 2A-D), which varied between 1-3 PDGFRα+ cells thick (SuppFigure 2E). Under this layer, we observed only a single layer of largely circumferential LMCs stained by Myh11 (SuppFigure 2B) sitting atop a single layer of CD31+ LECs (SuppFigure 2A). We also observed occasional PDGFRα+ cells or their extensions located in the sub-endothelial space (SuppFigure 2 E’, E”) positioned between the LECs and the LMCs.

Use of iCre-Mediated Recombination of Rosa26mT/mG to Delineate and Characterize Specific IALV Cell Types

Earlier version: After confirming the presence of Vimentin+, cKit+, and CD34+ PDGFRα+ positive cells within the mouse IALV, we sought to further investigate these cell populations by using constitutive and inducible Cre recombinase expressing mouse lines. IALVs from the constitutively active PdgfrαCre-ROSA26mTmG and Ng2Cre-ROSA26mTmG mice had GFP fluorescence in the majority of LMCs as well as in the fibroblast-shaped cells found within the IALV wall (Figure 4 A,B). While informative of expression of the LMC progenitor cells, neither constitutive Cre would be useful in delineating cell types. In contrast to the constitutively active PdgfrαCre, the tamoxifen inducible PdgfrαCreERTM line drove significant recombination in only the fibroblast-shaped cells previously stained with CD34 and PDGFRα but not in LMCs or LECs (Figure 4C). PdgfrβCreERT2, commonly used to label pericytes, drove recombination in both a minor population of the LMCs and the fibroblast-shaped cells. cKitCreERT2, which capably drives recombination in the ICC of the GI (Baker et al., 2016), drove recombination only in a small population of irregularly spaced large ovoid cells on the surface of the IALV (Figure 4E), although recombination in 1 or 2 LECs could occasionally be detected (not shown). Finally, Myh11CreERT2 drove recombination in nearly all LMCs which were largely circumferentially oriented with dendrite-like, cell-cell contacts visible between them and without significant GFP fluorescence in either LECs or the fibroblast-shaped CD34+ PDGFRα+ cell population (Figure 4F). Additionally, some LMCs maintained the bipolar shape but had secondary extensions forming a “Y” shape in which an adjacent LMC typically filled the inner void. A very minor population of recombined cells in the Myh11CreERT2-ROSA26mTmG IALVs were smaller and irregularly patterned with multiple fine axon-like projections or ruffled edges (Figure 4F).

Revised version: After confirming the presence of VIMENTIN+, cKIT+, and CD34+ PDGFRα+ positive cells within the mouse IALV, we sought to further investigate these cell populations by using constitutive and inducible Cre recombinase expressing mouse lines. IALVs from the constitutively active PdgfrαCre-ROSA26mTmG and Ng2Cre-ROSA26mTmG mice had GFP fluorescence in the majority of LMCs as well as in the fibroblast-shaped cells found within the IALV wall (Figure 4 A,B). While informative of expression of the LMC progenitor cells, neither constitutive Cre would be useful in delineating cell types. In contrast to the constitutively active PdgfrαCre, the tamoxifen inducible PdgfrαCreERTM line drove significant recombination in only the fibroblast-shaped cells previously stained for CD34 and PDGFRα but not in LMCs or LECs (Figure 4C). PdgfrβCreERT2, commonly used to label pericytes, drove recombination in both a minor population of the LMCs and the fibroblast-shaped cells. cKitCreERT2, which capably drives recombination in the ICCs of the GI tract (Baker et al., 2016), drove recombination only in a small population of irregularly-spaced, large ovoid cells on the surface of the IALV (Figure 4E), although recombination in 1 or 2 LECs could occasionally be detected (not shown). Finally, Myh11CreERT2drove recombination in nearly all LMCs which were largely circumferentially oriented with dendrite-like, cell-cell contacts visible between them and without significant GFP fluorescence in either LECs or the fibroblast-shaped CD34+ PDGFRα+ cell population (Figure 4F). Additionally, some LMCs maintained the bipolar shape but had secondary extensions forming a “Y” shape in which an adjacent LMC typically filled the inner void. A very minor population of recombined cells in the Myh11CreERT2-ROSA26mTmG IALVs were smaller and irregularly patterned with multiple fine axon-like projections or ruffled edges (Figure 4F).

iCre-ROSA26mTmG Labelling and Fidelity to Target Putative Pacemaker Cell Populations

Stitched montages of serial max projections of GFP and tdTomato signal from live IALVs isolated from PdgfrαCre-ROSA26mTmG (A), Ng2Cre-ROSA26mTmG (B), PdgfrαCreERTM-ROSA26mTmG (C), PdgfrβCreERT2-ROSA26mTmG (D), cKitCreERT2-ROSA26mTmG (E), and Myh11CreERT2- ROSA26mTmG (F). IALVs were digested into single cells and GFP+ cells were purified via FACS from Prox1-eGFP (G), Myh11CreERT2-ROSA26mTmG (H), PdgfrαCreERTM-ROSA26mTmG (I), and PdgfrβCreERT2-ROSA26mTmG (J) mice. Representative gels demonstrating RT-PCR products corresponding to the respective genes used in the promoter of each specific transgene employed to drive either eGFP or Cre mediated recombination of ROSA26mTmG from each GFP+ sorted population (K-N) to assess fidelity. Images are representative of IALVs from at least 3 separate mice. FACs and RT-PCR was repeated at least 3 times for each mouse.

To complement the morphological and cell density findings obtained with confocal microscopy, we digested IALVs from the iCre-ROSA26mTmG lines, and the Prox1-eGFP line as a control, into single cell suspensions and sorted the respective GFP+ populations (Figure 4G-J) for RT-PCR profiling (Figure 4K). We first focused on determining the molecular fidelity of the sorted cells based on the gene promoters used to drive each “iCre” model to discern cellular overlap. In agreement with the confocal images, sorted GFP+ cells from PdgfrβCreERT2-ROSA26mT/mG IALVs expressed Pdgfrβ but also Myh11 and Pdgfrα. In contrast, GFP-sorted cells from PdgfrαCreERTM IALVs expressed Pdgfrα and Pdgfrβ, but with no detectable expression of Myh11. GFP+ cells from sorted Myh11CreERT2-ROSA26mTmG IALVs had high expression for Myh11 as well as Pdgfrβ, but did not express Pdgfrα. IALVs from cKitCreERT2-ROSA26mTmG mice were not pursued for FACS due to the exceptionally sparse recombination observed along the IALV.

Earlier version: Characterization of the cellular constituents of the mouse IALVs by scRNAseq and FACs-RT-PCR

Revised version: Characterization of the cellular components of the mouse IALVs by scRNAseq and FACs-RT-PCR

Earlier version: The results from the immunofluorescence staining, ROSA26mTmG reporter imaging, and FACs-RT-PCR experiments suggested both LMCs and AdvCs express Pdgfrβ. To provide further clarity and detail to the cellular populations within the mouse cLV wall, and potential subsets within those broad cell types we performed scRNAseq on our isolated and cleaned inguinal axillary cLVs from male and female mice. The resulting UMAP projection (Figure 5A) revealed a host of cell types which had 3 main clusters which corresponded to LECs, LMCs and AdvCs (Figure 5A). We assessed the expression of genes that correspond to the markers from our earlier immunofluorescence staining as well as cell identification markers commonly used within the literature to identify each cell cluster (Figure 5B). Cell identity was confirmed by commonly used markers (Figure 5B) and the top differentially expressed genes (SuppFigure 3A). Dot plots for the LEC markers Prox1 (Figure 5C) and Flt4 (Figure 5D), LMC markers Myh11 (Figure 5E) and Cnn1 (Figure 5F), and the AdvCs markers Pdgfrα (Figure 5G) and Lumican (Figure 5H) were quite specific for labelling their respective cell clusters. Very few Kit (Figure 5I) expressing cells were observed in accordance with our imaging results. Pdgfrβ was observed in both LMC and AdvC clusters (Figure 5J) while the remaining cell clusters were of immune origin as they expressed the gene encoding the hematopoietic marker CD45 (Figure 5K). Notably, the previous genes suggested to identify LMCs in a previous scRNASeq study (Kenney et al., 2022), Dpt, Pi16, and Ackr3, were largely absent in LMCs and instead were expressed in a minor population of AdvCs (SuppFigure3B). We provide a further sub-clustering breakdown of the LECs (SuppFigure 4), LMCs (SuppFigure5), AdvCs (SuppFigure 6), and a detailed expression profile of the immune cell clusters (SuppFigure 7). Further assessment of the LEC subcluster included a putative lymphatic endothelial “up valve” cell population in sub cluster 8 which expressed high levels of Prox1, Cldn11, Itga9, Gja4, and Neo1 and “down valve” population in cluster 6 which expressed Clu, Adm, Gja4 and Lypd6 (SuppFigure 4C) which map well to a previous RNAseq dataset (Gonzalez-Loyola et al., 2021; Petkova et al., 2023; Yoon et al., 2024). The top differentially expressed genes in the putative down valve population in cluster 8 included Irx3, Neo1, Tub, Ano4, and Fxyd2 and we noted Cacna1e, Fgf14, and Irf1 in the up-valve cluster 6. Analysis of the LMC subclusters did not reveal any significant differences in known pacemaking associated genes Ano1 or Itpr1 expression. However, we provide an overview of the typical ion channel families expressed in LMCs in SuppFigure 5B-I. The AdvC cells could be further subclustered into multiple populations (SuppFigure 6A,C) with little evidence of LMC gene contamination as these cells lacked Myh11, Kcnma1, and Tagln though despite expression of Cacna1c, Ano1, and Cx45. Over 75% of AdvCs expressed Pdgfrα (SuppFigure 6) and 65% of the total AdvCs expressed both Pdgfrα and CD34. Our immunofluorescent colocalization of Pdgfrα and CD34 was also supported as 72% of Pdgfrα expressing AdvCs also co-expressed CD34 (SuppFigure 6D). The vast majority of AdvCs expressing Pdgfrβ (SuppFigure 6E) or Cspg4 (SuppFigure 6F) also expressed Pdgfrα. Expression of Ano1, Cx45, and Cacna1c, was also observed in some of the AdvCs and most of these cells also co-expressed Pdgfrα supporting further use of the PdgfrαCreERTM line (SuppFigure 6G-I).

Revised version: The results from the immunofluorescence staining, ROSA26mTmG reporter imaging, and FACs-RT-PCR experiments suggested that both LMCs and AdvCs can express Pdgfrβ. To provide further clarity and detail to the cellular populations within the mouse cLV wall and potential subsets within those broad cell types we performed scRNAseq on isolated and cleaned inguinal axillary cLVs from male and female mice. The resulting UMAP projection (Figure 5A) revealed a host of cell types which had 3 main clusters corresponding to LECs, LMCs and AdvCs (Figure 5A). We assessed the expression of genes that correspond to the markers from our earlier immunofluorescence staining as well as cell identification markers commonly used within the literature to identify each cell cluster (Figure 5B). Cell identity was confirmed by commonly used markers (Figure 5B) and the top differentially expressed genes (SuppFigure 3A). Dot plots for the LEC markers Prox1 (Figure 5C) and Flt4 (Figure 5D), LMC markers Myh11 (Figure 5E) and Cnn1 (Figure 5F), and the AdvCs markers Pdgfrα (Figure 5G) and Lumican (Figure 5H) were quite specific for labelling their respective cell clusters. Very few Kit (Figure 5I) expressing cells were observed in accordance with our imaging results. Pdgfrβ was observed in both LMC and AdvC clusters (Figure 5J) while the remaining cell clusters were of immune origin as they expressed the gene encoding the hematopoietic marker CD45 (Figure 5K). Notably, the previous genes suggested to identify LMCs in a previous scRNASeq study (Kenney et al., 2022), Dpt, Pi16, and Ackr3, were largely absent in LMCs and instead were expressed in a minor population of AdvCs (SuppFigure3B). We provide a further sub-clustering breakdown of the LECs (SuppFigure 4), LMCs (SuppFigure5), AdvCs (SuppFigure 6), and a detailed expression profile of the immune cell clusters (SuppFigure7). Further assessment of the LEC subcluster included a putative lymphatic endothelial “up valve” cell population in sub cluster 8 which expressed high levels of Prox1, Cldn11, Itga9, Gja4, and Neo1 and “down valve” population in cluster 6 which expressed Clu, Adm, Gja4 and Lypd6 (SuppFigure 4C) which mapped well to a previous RNAseq dataset (Gonzalez-Loyola et al., 2021; Petkova et al., 2023; Yoon et al., 2024). The top differentially expressed genes in the putative down valve population in cluster 8 included Irx3, Neo1, Tub, Ano4, and Fxyd2 and we noted Cacna1e, Fgf14, and Irf1 in the up-valve cluster 6. Analysis of the LMC subclusters did not reveal any significant differences in the expression of known pacemaking associated genes Ano1 or Itpr1 at our initial conditions of Log2FC of 0.5. However, we provide an overview of the typical ion channel families expressed in LMCs in SuppFigure 5B-I. The AdvC cells could be further subclustered into multiple populations (SuppFigure 6A,C) with little evidence of LMC gene contamination as these cells lacked Myh11, Kcnma1, and Tagln despite expression of Cacna1c, Ano1, and Cx45. Over 75% of AdvCs expressed Pdgfrα (SuppFigure 6) and 65% of the total AdvCs expressed both Pdgfrα and CD34. Our immunofluorescent colocalization of Pdgfrα and CD34 was also supported as 72% of Pdgfrα expressing AdvCs also co-expressed CD34 (SuppFigure 6D). The vast majority of AdvCs expressing Pdgfrβ (SuppFigure 6E) or Cspg4 (SuppFigure 6F) also expressed Pdgfrα. Expression of Ano1, Cx45, and Cacna1c, was also observed in some of the AdvCs and most of those cells also co-expressed Pdgfrα, supporting further use of the PdgfrαCreERTM line (SuppFigure 6G-I).

scRNAseq analysis of mouse IALVs from ROSA26mTmG mice.

IALVs were cleaned and isolated from 8 ROSA26mTmG mice and digested into a single cell suspension for scRNAseq analysis with the 10X platform. A) UMAP of the various cell populations that compromise the mouse IALV though some mammary epithelia contamination was present (populations 18,19). B) Heat map of commonly used genes for cell identification for each of the cell clusters. Dot plots to assess cell cluster expression of the genes shown in Figure 4 using a dot plot for the LEC markers Prox1 (C) and Flt4 (D, VEGFR3), LMC markers Myh11 (E) and caponin1 (F, Cnn1), fibroblast markers Pdgfrα (G) and Lum (H, Lumican), ICC marker Kit (I), the pericyte and smooth muscle precursor marker (Pdgfrβ), and the hematopoietic marker Ptprc (K, CD45).

Earlier version: While scRNASeq highlighted the depth of heterogeneity of the cellular composition of the mouse cLV, we wanted to validate the actual recombined cell populations from our iCre-ROSA26mTmG models. We profiled each iCre driven recombination of ROSAmTmG via FACs-purified cells and RT-PCR for common markers for endothelial cells, muscle cells, and pericytes. Nos3 (eNOS) expression was observed only in the Prox1-eGFP sorted cells, and LECs also expressed Vim, Mcam, and had weak but detectable signal for CD34 (Figure 6A). Myh11CreERT2 sorted cells showed expression of smooth muscle actin (Acta2), the alpha subunit of the L-type voltage gated Ca2+ channel Cacna1c (Cav1.2), Desmin (Des), Mcam, and Vimentin (Vim, Figure 6B). In addition to the genes expressed under Myh11CreERT2 recombination, Cdh5, CD34, and Cspg4 (Ng2) were detected in cells sorted from PdgfrβCreERT2 IALVs (Figure 6C). As expected, the GFP+ cells sorted from PdgfrαCreERTM IALVs expressed mRNA for CD34, weak signal for Cspg4, and Vimentin, but not Desmin, Acta2, nor the pericyte marker Mcam (Figure 6D). The Cacna1c was expressed in cells FACS purified from both PdgfrβCreERT2 and Myh11CreERT2 IALVs and sorted cells from PdgfrαCreERTM IALVs without any evidence that Myh11 expressing muscle cells contaminated the latter. These findings confirmed the separate cell populations achieved with PdgfrαCreERTM and Myh11CreERT2 mediated recombination, at least as it pertains to the ROSA26mTmG reporter. These findings were largely validated by our scRNASeq dataset. Cdh5 (Figure 6E) and Nos3 (Figure 6F) were almost exclusively expressed in our LEC clusters while Acta2 (Figure 6G) was highly expressed in the LMC cluster. We also observed that Cacna1c was highly expressed in the LMCs (Figure 6H) and some AdvCs. Cd34 was highly expressed in the AdvCs but was also observed in LECs (Figure 6I) although we did not observe a signal in our earlier immunofluorescence staining (Figure 3). Cspg4 was observed in a minor population of AdvCs (Figure 6J). The intermediate filament Vim (Figure 6K) was ubiquitously across all clusters expressed but Des was primarily expressed in LMCs and some subsets of AdvCs (Figure 6K,L). The endothelial and pericyte marker Mcam (also referred to as CD146) was expressed in LECs and LMCs but was largely absent in AdvCs (Figure 6M). We followed up the identification of Cav1.2 expression in the PdgfrαCreERTM sorted cell population by assessing the expression of other genes involved in either electrical conduction (Cx45) (Figure 6N) or pacemaking (Ano1) (Figure 6O) of IALVs. Expression of Ano1 and Cx45 was observed in PdgfrαCreERTM ROSA26mtmG FACS-purified cells (Figure 6P).

Revised version: While scRNASeq highlighted the depth of heterogeneity of the cellular composition of the mouse cLV, we wanted to validate the actual recombined cell populations from the iCre-ROSA26mTmG models. We profiled each iCre driven recombination of ROSAmTmG via FACs-purified cells and RT-PCR for common markers for endothelial cells, muscle cells, and pericytes. Nos3 (eNOS) expression was observed only in the Prox1-eGFP sorted cells, and LECs also expressed Vim, Mcam, and had weak but detectable signal for CD34 (Figure 6A). Myh11CreERT2 sorted cells showed expression of smooth muscle actin (Acta2), the alpha subunit of the L-type voltage gated Ca2+ channel Cacna1c (Cav1.2), Desmin (Des), Mcam, and Vimentin (Vim, Figure 6B). In addition to the genes expressed under Myh11CreERT2 recombination, Cdh5, CD34, and Cspg4 (Ng2) were detected in cells sorted from PdgfrβCreERT2 IALVs (Figure 6C). As expected, the GFP+ cells sorted from PdgfrαCreERTM IALVs expressed mRNA for CD34, weak signal for Cspg4, and Vimentin, but not Desmin, Acta2, nor the pericyte marker Mcam (Figure 6D). Cacna1c was expressed in cells FACS purified from both PdgfrβCreERT2 and Myh11CreERT2 IALVs and sorted cells from PdgfrαCreERTM IALVs without any evidence that Myh11 expressing muscle cells contaminated the latter. These findings confirmed the separate cell populations achieved with PdgfrαCreERTMand Myh11CreERT2 mediated recombination, at least as it pertains to the ROSA26mTmG reporter. These findings were largely validated by our scRNASeq dataset. Cdh5 (Figure 6E) and Nos3 (Figure 6F) were almost exclusively expressed in the LEC clusters while Acta2 (Figure 6G) was highly expressed in the LMC cluster. We also observed that Cacna1c was highly expressed in the LMCs (Figure 6H) and some AdvCs. Cd34 was widely expressed in AdvCs matching our immunofluorescence data. Cd34 expression was also seen in LECs (Figure 6I) although we did not observe a signal in LECs in our earlier immunofluorescence staining (Figure 3). Cspg4 was observed in a minor population of AdvCs (Figure 6J). The intermediate filament Vim (Figure 6K) was ubiquitously expressed across all clusters expressed but Des was primarily expressed in LMCs and some subsets of AdvCs (Figure 6K, L). The endothelial and pericyte marker Mcam (also referred to as CD146) was expressed in LECs and LMCs but was largely absent in AdvCs (Figure 6M). We followed up the identification of Cav1.2 expression in the PdgfrαCreERTMsorted cell population by assessing the expression of other genes involved in either electrical conduction (Cx45) (Figure 6N) or pacemaking (Ano1) (Figure 6O) of IALVs. Expression of Ano1 and Cx45 was observed in PdgfrαCreERTMROSA26mtmG FACS-purified cells (Figure 6P).

RT-PCR Profiling of FACs Purified Cells from iCre-ROSA26mTmG

Expanded RT-PCR profiling of genes to discriminate LECs, LMCs, and other cell types in our GFP+ sorted cells from Prox1-eGFP (A), Myh11CreERT2-ROSA26mTmG (B), PdgfrβCreERT2-ROSA26mTmG (C), and PdgfrαCreERTM-ROSA26mTmG (D). Dot plots for the genes assessed in A-D in our IALV scRNAseq analysis confirmed those results. In addition to a population of AdvCs expressing Cacna1c, we also noted expression of Cx45 (N) which was also observed in LECs) and Ano1 (O) in the AdvC clusters. We confirmed this expression using GFP+ cells sorted from PdgfrαCreERTM-ROSA26mTmG IALVs for RT-PCR (P) and ruled out hematopoietic or LEC contamination. All RT-PCRs were performed 2-4 times for each gene over each sorted cell population collected from different mice.

Earlier version: Inducible Deletion of Either Cav1.2, Ano1, and Cx45 with PdgfrαCreERTM Did Not Affect cLV Pacemaking

Revised version: Inducible Deletion of Either Cav1.2, Ano1, or Cx45 with PdgfrαCreERTM Did Not Affect cLV Pacemaking

Earlier version: The expression of the genes critically involved in cLV function—Cav1.2, Ano1, and Cx45—in the PdgfrαCreERTM-ROSA26mTmG purified cells and scRNAseq data prompted us to generate PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice for contractile tests. We isolated popliteal cLVs and tested their pacemaker and contractile function in response to a physiological pressures range of 0.5-10 cmH2O, under normal conditions. However, we did not detect any significant differences in pacemaking nor contractile function as assessed by contraction frequency, ejection fraction, and vessel tone in popliteal cLVs studied from PdgfrαCreERTM-Ano1fl/fl mice (Figure 7A-C) or PdgfrαCreERTM-Cx45fl/fl mice (Figure7D-F). There was no difference in contraction frequency in cLVs from PdgfrαCreERTM-Cav1.2fl/fl mice, however, we noted a mild but statistically significant increase in ejection fraction at the lowest pressure, 0.5 cmH2O (Figure 7H). Additionally, vessels isolated from PdgfrαCreERTM-Cav1.2fl/fl mice also had a statistically significant increase in vessel tone (Figure 7I) noted at the 2-way level although we did not resolve significance at any specific pressure with this sample. No difference in normalized contraction amplitude, fractional pump flow, or diastolic diameters were observed (SuppFigure 8). In total, despite the presence of transcript for these critical genes in Pdgfrα+ cells, PdgfrαCreERTM mediated deletion of Cx45, Cav1.2 or Ano1 failed to recapitulate previous reports of the significant contractile defects using the Myh11CreERT2 line (Castorena-Gonzalez et al., 2018b; Zawieja et al., 2019; To et al., 2020; Davis et al., 2022).

Revised version: The expression of the genes critically involved in cLV function—Cav1.2, Ano1, and Cx45—in the PdgfrαCreERTM-ROSA26mTmG purified cells and scRNAseq data prompted us to generate PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice for contractile tests. We isolated popliteal cLVs and tested their pacemaker and contractile function in response to a physiological pressure range of 0.5-10 cmH2O, under normal conditions. However, we did not detect any significant differences in pacemaking or contractile function as assessed by contraction frequency, ejection fraction, and vessel tone in popliteal cLVs studied from PdgfrαCreERTM-Ano1fl/fl mice (Figure 7A-C) or PdgfrαCreERTM-Cx45fl/fl mice (Figure7D-F). There was no difference in contraction frequency of cLVs from PdgfrαCreERTM-Cav1.2fl/fl mice compared to floxed control mice, however, we noted a mild but statistically significant increase in ejection fraction at the lowest pressure, 0.5 cmH2O (Figure 7H). Additionally, vessels isolated from PdgfrαCreERTM-Cav1.2fl/fl mice also had a statistically significant increase in vessel tone (Figure 7I) noted at the 2-way level although we did not resolve significance at any specific pressure with this sample. No differences in normalized contraction amplitude, fractional pump flow, or diastolic diameter were observed (SuppFigure 8). In total, despite the presence of transcript for these critical genes in Pdgfrα+ cells, PdgfrαCreERTM mediated deletion of Cx45, Cav1.2 or Ano1 failed to recapitulate previous reports of the significant contractile defects using the Myh11CreERT2 line to delete the same genes (Castorena-Gonzalez et al., 2018b; Zawieja et al., 2019; To et al., 2020; Davis et al., 2022).

Isobaric contractile Assessment of popliteal cLV from PdgfrαCreERTM driven deletion of Ano1, CX45, and CaV1.2

Earlier version: Summary of the contractile parameters recorded from popliteal cLVs in PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl mice, PdgfrαCreERTM-Cav1.2fl/fl mice. Contraction frequency (A, D, G), ejection fraction (B, E, H), and vessel tone (C, F, I) were assessed. No statically significant differences observed in cLVs isolated from PdgfrαCreERTM-Ano1fl/fl and PdgfrαCreERTM-Cx45fl/fl mice across these three parameters. Mean and SEM shown, n=6 popliteal vessels from 3 mice PdgfrαCreERTM-Ano1fl/fl mice and n=10 popliteal vessels from 6 mice Ano1fl/fl mice. Mean and SEM shown, n=5 popliteal vessels from 3 mice PdgfrαCreERTM-CX45fl/fl mice and n=8 popliteal vessels from 11 mice CX45fl/fl mice. Mean and SEM shown, n=6 popliteal vessels from 3 mice PdgfrαCreERTM-Cav1.2fl/fl mice and n=9 popliteal vessels from 20 mice Cav1.2fl/fl mice. The contractile data from control Cav1.2fl/fl vessels was previously published but was separated by sex (Davis et al., 2022) while they are combined here. * Denotes significance at p <0.05 which 0.10 > p >0.05 are reported as text. Normalized contraction amplitude, fractional pump flow, end diastolic diameter can be found in Supp Figure 8.

Revised version: Summary of the contractile parameters recorded from popliteal cLVs in PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl mice, PdgfrαCreERTM-Cav1.2fl/fl mice. Contraction frequency (A, D, G), ejection fraction (B, E, H), and vessel tone (C, F, I) were assessed. No statically significant differences observed in cLVs isolated from PdgfrαCreERTM-Ano1fl/fl and PdgfrαCreERTM-Cx45fl/fl mice across these three parameters. Mean and SEM shown, n=6 popliteal vessels from 3 mice PdgfrαCreERTM-Ano1fl/fl mice and n=10 popliteal vessels from 6 mice Ano1fl/fl mice. Mean and SEM shown, n=5 popliteal vessels from 3 mice PdgfrαCreERTM-CX45fl/fl mice and n=8 popliteal vessels from 11 mice CX45fl/fl mice. Mean and SEM shown, n=6 popliteal vessels from 3 mice PdgfrαCreERTM-Cav1.2fl/fl mice and n=9 popliteal vessels from 20 mice Cav1.2fl/fl mice. The contractile data from control Cav1.2fl/fl vessels was previously published but was separated by sex (Davis et al., 2022) while they are combined here. * Denotes significance at p <0.05 which 0.10 > p >0.05 are reported as text. Normalized contraction amplitude, fractional pump flow, end diastolic diameter can be found in SuppFigure 8.

PDGFRα+ Adventitial Fibroblasts Express Markers Associated with Multipotency

Earlier version: Despite the lack of cLV pacemaking deficits in the PdgfrαCreERTM genetic knockout lines, we were curious to discern further insight into the role or function of the PDGFRα+ CD34+ cells since they comprise a significant portion of the lymphatic cLV wall. We performed RT-PCR on FACS purified cells from Prox1-eGFP, Myh11CreERT2-ROSA26mTmG, and PdgfrαCreERTM-ROSA26mTmG IALVs for markers multipotency including Krüppel-like factor 4 (Klf4), stem cell antigen 1 (Sca1, also referred to as Ly6a), and Gli1 with CD34 and Pdgfrα used to assess purity. Recombined (GFP+) cells from Myh11CreERT2-ROSA26mTmG had weak expression of Klf4 and Gli1 but were negative for Ly6a (SuppFigure 9A). PdgfrαCreERTM recombined cells strongly expressed Klf4, Ly6a, and Gli1 (SuppFigure 9A). LECs sorted from Prox1-eGFP IALVs were positive for Klf4, weak for Ly6a, and positive for CD34 but negative for Gli1 and PDGFRα (SuppFigure 9B). The unrecombined population (tdTomato+) cells in the Myh11CreERT2-ROSA26mTmG IALVs (SuppFigure 9B) showed expression for all the markers as expected. PdgfrαCreERTM recombined cells also expressed the mesenchymal stromal cell markers CD29, CD105, and CD44 (SuppFigure 9C, positive control in 9D). However, expression of these genes was not homogenous across all the AdvCs population based on our scRNAseq analysis (SuppFigure 9E-J). We performed immunofluorescence staining for one of these multipotent markers, Ly6a (SuppFigure 9K) in the adventitial cells with PDGFRα (SuppFigure 9L) and counter staining for LMCs with MYH11 (SuppFigure 9M). The morphology and staining pattern of Sca1 overlapped significantly with PDGFRα staining and not Myh11 staining (SuppFigure 9N, Supplemental Movie 6).

Revised version: Despite the lack of cLV pacemaking deficits in the PdgfrαCreERTMgenetic knockout lines, we were curious to discern further insight into the role or function of the PDGFRα+ CD34+ cells since they comprise a significant portion of the lymphatic cLV wall. We performed RT-PCR on FACS purified cells from Prox1-eGFP, Myh11CreERT2-ROSA26mTmG, and PdgfrαCreERTM-ROSA26mTmG IALVs for multipotency markers including Krüppel-like factor 4 (Klf4), stem cell antigen 1 (Sca1, also referred to as Ly6a), and Gli1, with CD34 and Pdgfrα used to assess purity. Recombined (GFP+) cells from Myh11CreERT2-ROSA26mTmG had weak expression of Klf4 and Gli1 but were negative for Ly6a (SuppFigure 9A). PdgfrαCreERTM recombined cells strongly expressed Klf4, Ly6a, and Gli1 (SuppFigure 9A). LECs sorted from Prox1-eGFP IALVs were positive for Klf4, weak for Ly6a, and positive for CD34 but negative for Gli1 and PDGFRα (SuppFigure 9B). The unrecombined population (tdTomato+) cells in the Myh11CreERT2-ROSA26mTmG IALVs (SuppFigure 9B) showed expression for all the markers as expected. PdgfrαCreERTM recombined cells also expressed the mesenchymal stromal cell markers CD29, CD105, and CD44 (SuppFigure 9C, positive control in 9D). However, expression of these genes was not homogenous across all the AdvCs population based on our scRNAseq analysis (SuppFigure 9E-J). We performed immunofluorescence staining for one of these multipotent markers, Ly6a (SuppFigure 9K) in the adventitial cells with PDGFRα (SuppFigure 9L) and counter staining for LMCs with MYH11 (SuppFigure 9M). The morphology and staining pattern of Sca1 overlapped significantly with PDGFRα staining and not MYH11 staining (SuppFigure 9N, Supplemental Movie 6).

Earlier version: Optogenetic Stimulation of iCre-driven Channel Rhodopsin 2 to Induce Test Light-Stimulated Depolarization Induced Lymphatic Contraction

Revised version: Optogenetic Stimulation of iCre-driven Channel Rhodopsin 2

Earlier version: We next used optogenetic methods to test whether the cell populations recombined by either cKitCreERT2, PdgfrαCreERTM, or Myh11CreERT2 could elicit a coordinated contraction. The ChR2-tdTomato construct appeared more sensitive to recombination than ROSA26mTmG, in some cases resulting in LMC expression of ChR2-tdTomato in PdgfrαCreERTM and CkitCreERT2 popliteal cLVs based on morphology. Care was taken to image each vessel for tdTomato (Figure 8A,C,E) prior to stimulation at its respective sites under brightfield conditions for diameter tracking (Figure 8B,D,F) to ensure fidelity of the cell types and morphologies observed in Figure 3 and Figure 4. As with ROSA26mTmG, CkitCreERT2 drove the ChR2-tdTomato expression primarily in large ovoid cells found on the adventitia of the vessel. Cells were stimulated by positioning an optical laser fiber (tip diameter 2-3 mm) near a ChR2+ cell, with an illumination field of 10-50 mm. Localized photo-stimulation of these cells did not initiate coordinated contractions (Figure 8G-J,S). Similarly, photo-stimulation of ChR2-tdTomato expressing cells driven by PdgfrαCreERTM failed to initiate a coordinated contraction (Figure 8K-N, T). In contrast, localized photo-stimulation of LMCs, using Myh11CreERT2 to express Chr2-tdTomato, resulted in a propagated contraction in the popliteal vessel (Figure 8O-R, U). In total, only 3.25% of photo-stimulation events for cKitCreERT2-ChR2-TdTomato and 3.03% of photo-stimulation events for PdgfrαCreERTM-ChR2-tdTomato were associated with a contraction, while 88.5% of photo-stimulation events for Myh11CreERT2-ChR2-tdTomato induced contractions (Figure 8V). The optogenetic triggering of contractions observed in PdgfrαCreERTM-ChR2-tdTomato and cKitCreERT2-ChR2-TdTomato vessels is likely due to the happenstance of spontaneous contractions occurring during the time and proximity of optogenetic stimulation (see Methods). As a control, we also used non-induced (no tamoxifen) Myh11CreERT2-ChR2-tdTomato cLVs and contractions were associated with only 7% of photo-stimulation events, in line with the PdgfrαCreERTM and cKitCreERT2 results (Figure 8V). As mast cells are not ascribed any tissue specific pacemaking behavior, these similar low percentages observed between these three groups are suggestive of random coincidence. Brightfield videos of the photo-stimulation and representative traces for cKitCreERT2-ChR2-TdTomato, PdgfrαCreERTM-ChR2-tdTomato, Myh11CreERT2-ChR2-tdTomato are provided in Supplemental Movies 7-9.

Revised version: We next used optogenetic methods to test whether the cell populations recombined by either cKitCreERT2, PdgfrαCreERTM, or Myh11CreERT2 could elicit a coordinated contraction. The ChR2-tdTomato construct appeared more sensitive to recombination than ROSA26mTmG, in some cases resulting in LMC expression of ChR2-tdTomato in PdgfrαCreERTMand cKitCreERT2 popliteal cLVs based on cell morphology. Care was taken to image each vessel for tdTomato (Figure 8A,C,E) prior to stimulation at its respective sites under brightfield conditions for diameter tracking (Figure 8B,D,F) to ensure fidelity of the cell types and morphologies observed in Figure 3 and Figure 4. As with ROSA26mTmG, cKitCreERT2 drove the ChR2-tdTomato expression primarily in large ovoid cells found on the adventitia of the vessel. Cells were stimulated by positioning an optical laser fiber (tip diameter 2-3 μm) near a ChR2+ cell, with an illumination field of 10-50 μm. Localized photo-stimulation of these cells did not initiate coordinated contractions (Figure 8G-J,S). Similarly, photo-stimulation of ChR2-tdTomato expressing cells driven by PdgfrαCreERTM failed to initiate a coordinated contraction (Figure 8K-N, T). In contrast, localized photo-stimulation of LMCs, using Myh11CreERT2 to express Chr2-tdTomato, resulted in a propagated contraction in the popliteal vessel (Figure 8O-R, U). In total, only 3.25% of photo-stimulation events for cKitCreERT2-ChR2-TdTomato and 3.03% of photo-stimulation events for PdgfrαCreERTM-ChR2-tdTomato were associated with a contraction, while 88.5% of photo-stimulation events for Myh11CreERT2-ChR2-tdTomato induced contractions (Figure 8V). The optogenetic triggering of contractions observed in PdgfrαCreERTM-ChR2-tdTomato and cKitCreERT2-ChR2-TdTomato vessels is likely due to the happenstance of spontaneous contractions occurring during the time and proximity of optogenetic stimulation (see Methods). As a control, we also used non-induced (no tamoxifen) Myh11CreERT2-ChR2-tdTomato cLVs and contractions were associated with only 7% of photo-stimulation events, in line with the PdgfrαCreERTM and cKitCreERT2 results (Figure 8V). As mast cells are not ascribed any tissue specific pacemaking behavior, these similar low percentages observed between these three groups are suggestive of random coincidence. Brightfield videos of the photo-stimulation and representative traces for cKitCreERT2-ChR2-TdTomato, PdgfrαCreERTM-ChR2-tdTomato, Myh11CreERT2-ChR2-tdTomato are provided in Supplemental Movies 7-9.

ChR2-Mediated Depolarization Only in LMCs Triggers Contraction

Representative max projections of tdTomato-ChR2 signal in popliteal cLVs isolated from cKitCreERT2- ChR2-tdTomato (A), PdgfrαCreERTM-ChR2-tdTomato (C), and Myh11CreERT2- ChR2-tdTomato (E) with their corresponding brightfield image (B, D, F) respectively. Time-lapse brightfield images every 0.5 s starting at stimulation t=0 for cKitCreERT2-ChR2-tdTomato (G-J), PdgfrαCreERTM-ChR2-tdTomato (K- N), and Myh11CreERT2- ChR2-tdTomato (O-R). The I bar denotes the inner diameter at t=0 over time and white asterisks denote the contraction. Representative diameter trace for the popliteal cLV demonstrate spontaneous contractions with the dotted boxes indicating the optical stimulation event in the respective brightfield images of the time lapse images. Isolated cLVs from cKitCreERT2-ChR2-tdTomato (S), PdgfrαCreERTM-ChR2-tdTomato (T), and Myh11CreERT2- ChR2-tdTomato (U) were stimulated with light pulses (red dashed lines) and the summation of contraction triggering for each genotype (V). Mean and SEM are shown, **** denotes p<0.0001. Contraction recorded from at least 6 popliteal cLVs from 3 mice per genotype.

Confocal Ca2+ Imaging of GCaMP6f Expression Driven by cKitCreERT2, PdgfrαCreERTM, and Myh11CreERT2 Over the Lymphatic Contraction Cycle

Earlier version: Subcellular calcium transients are observed in many pacemaker cells. We imaged IALVs from cKitCreERT2-GCaMp6f mice, which primarily resulted in expression of GCaMp6f in the large ovoid cells in the adventitia (Figure 9A), although we occasionally observed GCaMP6f expression in both LEC and LMCs (Figure 9A) as depicted in the maximum projection of the acquisition period (Supplemental Movie 10) and the spatio-temporal maps (STMS). The aberrant expressions of GCaMP6f in cells that demonstrated the typical cobblestone morphology of LECs or the circumferential LMCs that exhibited Ca2+ flashes and diastolic Ca2+ transients (Figure 9D,E green arrows) prior to contraction were not included in the cKitCreERT2-GCaMp6f analysis. Of 39 cKitCreERT2-GCaMp6f cells analyzed, only 1 cKitCreERT2-GCaMP6f cell exhibited a spontaneous Ca2+ transient during the recording period (Figure 9B,C Cell 7). However, the Ca2+ transient in that cell did not align temporally with the “Ca2+ flash” of the LMC with incidental GCaMp6f expression (Figure 9C,D). Despite the lack of Ca2+ transients under the baseline conditions throughout the IALV contraction cycle, many cKitCreERT2-GCaMP6f cells exhibited a robust and prolonged Ca2+ event in response to stimulation with the mast cell activator compound 48-80 (Figure 9F, G, H). Notably, the Ca2+ events in the ovoid cells elicited by administration of compound 48-80 did not acutely alter the LMC Ca2+ activity (Figure 9I,J). Like cKitCreERT2-GCaMp6f, the majority of PDFRnCreERTM-GCaMP6f expressing cells also largely lacked Ca2+ transients and resulted in incidental LMC GCaMP6f expression (Figure 10B, Supplemental Movie 11). Some cells exhibited high basal Ca2+ levels (Figure 10A,D) sustained throughout the recording, but without oscillations (Figure 10B,C). In contrast, spurious GCaMP6f expression in a circumferentially oriented LMC displayed Ca2+ flashes associated with contraction (Figure 10B,C). Of the 21 PDGFRα-GCaMP6f cells assessed, only 3 exhibited Ca2+ transients that were singular events with limited spatial spread within the 20 sec imaging period (Figure 10E,F). The lack of either global or consistent Ca2+ transients within either cKitCreERT2-GCaMP6f or PdgfrαCreERTM-GCaMP6f IALVs was in stark contrast to Ca2+ imaging of Myh11CreERT2-GCaMP6f IALVs. Myh11CreERT2 drove GCaMp6f expression in the circumferential LMCs (Figure 11A), which had global and nearly synchronous Ca2+ flashes in 100% of the analyzed cells (Figure 11B, C). Additionally, non-synchronous stochastic and localized Ca2+ transients during diastole were commonly observed in the LMCs (Figure 11D, E, Supplemental Movie 12). Many LMCs exhibited Ca2+ transients during each diastolic period while other LMCs displayed few Ca2+ transients or lacked diastolic Ca2+ transients during the recording period (Figure 11B). In aggregate, of the 39 cKitCreERT2-GCAMP6f cells only 1 displayed a Ca2+ transient during recording, 3 of 21 PdgfrαCreERTM-GCaMP6f cells, while 20 of 43 LMCs displayed at least one diastolic transient apart from 43 of 43 LMCs with global flashes.

Revised version: Subcellular calcium transients are observed in many pacemaker cells. We imaged IALVs from cKitCreERT2-GCaMP6f mice, which primarily resulted in expression of GCaMP6f in the large ovoid cells in the adventitia (Figure 9A), although we occasionally observed GCaMP6f expression in both LEC and LMCs (Figure 9A) as depicted in the maximum projection of the acquisition period (Supplemental Movie 10) and the spatio-temporal maps (STMS). The aberrant expressions of GCaMP6f in cells that demonstrated the typical cobblestone morphology of LECs or the circumferential LMCs that exhibited Ca2+ flashes and diastolic Ca2+ transients (Figure 9D,E green arrows) prior to contraction were not included in the cKitCreERT2-GCaMP6f analysis. Of 39 cKitCreERT2-GCaMP6f cells analyzed, only 1 cKitCreERT2-GCaMP6f cell exhibited a spontaneous Ca2+ transient during the recording period (Figure 9B,C Cell 7). However, the Ca2+ transient in that cell did not align temporally with the “Ca2+ flash” of the LMC with incidental GCaMP6f expression (Figure 9C,D). Despite the lack of Ca2+ transients under the baseline conditions throughout the IALV contraction cycle, many cKitCreERT2-GCaMP6f cells exhibited a robust and prolonged Ca2+ event in response to stimulation with the mast cell activator compound 48-80 (Figure 9F, G, H). Notably, the Ca2+ events in the ovoid cells elicited by administration of compound 48-80 did not acutely alter the LMC Ca2+ activity (Figure 9I,J). Similarly, the majority of PdgfrαCreERTM-GCaMP6f expressing cells also largely lacked Ca2+ transients and resulted in incidental LMC GCaMP6f expression (Figure 10B, Supplemental Movie 11). Some cells exhibited high basal Ca2+ levels (Figure 10A,D) sustained throughout the recording, but without oscillations (Figure 10B,C). In contrast, spurious GCaMP6f expression in a circumferentially oriented LMC displayed Ca2+ flashes associated with contraction (Figure 10B,C). Of the 21 PdgfrαCreERTM -GCaMP6f cells assessed, only 3 exhibited Ca2+ transients and those were singular events contained within a single cell within the 20 sec imaging period (Figure 10E,F). The lack of either global or consistent Ca2+ transients within either cKitCreERT2-GCaMP6f or PdgfrαCreERTM-GCaMP6f IALVs was in stark contrast to Ca2+ imaging of Myh11CreERT2-GCaMP6f IALVs. Myh11CreERT2 drove GCaMP6f expression in nearly all circumferential LMCs (Figure 11A), which exhibited global and nearly synchronous Ca2+ flashes in 100% of the analyzed cells (Figure 11B, C). Additionally, non-synchronous stochastic and localized Ca2+ transients were commonly observed in the LMCs during diastole (Figure 11D, E, Supplemental Movie 12). Many LMCs exhibited Ca2+ transients during each diastolic period while other LMCs displayed few Ca2+ transients or lacked diastolic Ca2+ transients during the recording period (Figure 11B). In aggregate, only 1 of 39 cKitCreERT2-GCaMP6f cells and 3 of 21 PdgfrαCreERTM-GCaMP6f cells displayed a Ca2+ transient during recording, while 20 of 43 LMCs displayed at least one diastolic transient apart from 43 of 43 LMCs with global Ca2+ flashes.

cKitCreERT2 Drives GCaMP6f Expression Primarily in Mast Cells in Mouse IALVs

Representative max projection of GCaMP6f signal over time in an IALV isolated from a cKitCreERT2- GCaMP6f mouse with ROI indicated around individual cells, primarily large ovoid cells, but also including a circumferential LMC (Cell10) and a horizontal LEC (Cell 11). Of cells 1-9, only cell 7 had any Ca2+ activity (red arrows) during the recording time as indicated by the STMs from each ROI (B) and their normalized F/F0 plots in (C). In contrast, the LMC in ROI 10 had both rhythmic global Ca2+ events (D) that spanned the cell axis (vertical axis) in the STM (E) in addition to localized Ca2+ events intervening the time between global events (green arrows). Representative max projection of GCaMP6f signal over time after stimulation with C48-80 (F) with many large ovoid cells displaying long lasting global Ca2+ events (G, H) while not immediately affecting the LMC Ca2+ dynamics (I, J).

Lack of coordinated Ca2+ Activity Across Contraction Cycle in PDGFRα Cells

Representative max projections of GCaMP6f signal over time in an IALVs isolated from PdgfrαCreERTM- GCaMP6f mice (A, D). ROIs were made around cells and GCaMP6f recorded over time to generate the corresponding STMs (B, E) for each cell and plots (C, F) respectively. Once again, incidental recombination occurred in a LMC which displayed rhythmic Ca2+ flashes (C) while the slight undulation in the other cells is due to movement artifact (B). Red arrows indicate the limited local Ca2+ activity observed in two cells from a PdgfrαCreERTM-GCaMP6f IALV.

Heterogeneous Diastolic Ca2+ Transient Activity in LMCs

Earlier version: Representative max projections of GCaMP6f signal over time in an IALVs isolated from Myh11CreERT2-GCaMP6f mice (A). LMCs were outlined with ROIs to assess GCaMp6F signal over time. Rhythmic global flashes (B) were entrained across all the LMCs in the FOV (C) with many cells exhibiting diastolic Ca2+ release events. Cells exhibiting at least one diastolic Ca2+ event, within the context of our focal plane constraints, over the recorded time were denoted by the red asterisks. The plot in (D) magnifies the first diastolic period, seconds 1-3 of C to assist in visualizing the lack of coordination of the diastolic events. (D) Max projection of the pseudo-linescan analysis across the axis of the vessel to highlight diastolic Ca2+ transients in all cells in the field of view and their lack of coordination across the cells (x-axis). The white dotted box shows the first diastolic period plotted in (D).

Revised version: Representative max projections of GCaMP6f signal over time in an IALVs isolated from Myh11CreERT2-GCaMP6f mice (A). LMCs were outlined with ROIs to assess GCaMp6F signal over time. Rhythmic global flashes (B) were entrained across all the LMCs in the FOV (C) with many cells exhibiting diastolic Ca2+ release events. Cells exhibiting at least one diastolic Ca2+ event, within the context of our focal plane constraints, over the recorded time were denoted by the red asterisks. The plot in (D) magnifies the first diastolic period, seconds 1-3 of C to assist in visualizing the lack of coordination of the diastolic events. (E) Max projection of the pseudo-linescan analysis across the axis of the vessel to highlight diastolic Ca2+ transients in all cells in the field of view and their lack of coordination across the cells (x-axis). The white dotted box shows the first diastolic period plotted in (D).

Pressure Dependency of Subcellular Ca2+ Transients in LMCs

Earlier version: We next sought to test whether diastolic Ca2+ transients were pressure-dependent, given that cLVs exhibit pressure dependent chronotropy (Zawieja et al., 2019). GCaMP6f expressing LMCs were studied at intraluminal pressures of 0.5 -5 cmH2O in the presence of nifedipine, which blocks the Ca2+ flashes but not local Ca2+ transients (Figure 12A). As intra-luminal pressure was increased, there was a marked increase in the occurrence of Ca2+ transients (Figure 12B, Supplemental Movies 13-15). We converted these calcium transients into particles (PTCLs) for further analysis as previously described. We generated activity maps of Ca2+ PTCL activity (Figure 12C) and determined PTCL area (Figure 12D) and frequency at each pressure (Figure 12E). The maps show that as pressure increased, the activity of PTCLs across the vessel also increased (as evident by the increase in PTCL area activation). Across 11 experiments, the area of the field of view activated by PTCLs/frame increased from 73.2 ± 17.7 µM 2/frame at 0.5 cmH20 to 108.6 ± 20.5 µM 2/frame at 2 cm H20 and further enhanced to 139.2 ± 26.9 µM 2/frame at 5 cm H2O (Figure 12F). The number of PTCLs per frame also increased with pressure, from 2.9 ± 0.4 at 0.5 cmH20 to 4.1 ± 0.5 and 5.2 ± 0.6 PTCL/frame at 2 and 5 cmH20 respectively (Figure 12G).

Revised version: We next sought to test whether diastolic Ca2+ transients were pressure-dependent, given that cLVs exhibit pressure-dependent chronotropy (Zawieja et al., 2019). GCaMP6f expressing LMCs were studied at intraluminal pressures of 0.5 -5 cmH2O in the presence of nifedipine, which blocked the Ca2+ flashes but not local Ca2+ transients (Figure 12A). As intra-luminal pressure was increased, there was a marked increase in the occurrence of Ca2+ transients (Figure 12B, Supplemental Movies 13-15). These calcium transients were converted into particles (PTCLs) for further analysis as previously described (Drumm et al., 2019a). Activity maps of Ca2+ PTCL activity were generated (Figure 12C) and PTCL area (Figure 12D) and frequency were determined at each pressure (Figure 12E). The maps show that as pressure increased, the area of the LMC layer displaying a Ca2+ transient increased (as evident by the increase in PTCL area) as did the distribution of Ca2+ PTCLs across the LMC layer (Figure 12C). Across 11 experiments, the area of the field of view activated by PTCLs/frame increased from 73.2 ± 17.7 µm 2/frame at 0.5 cmH20 to 108.6 ± 20.5 µm 2/frame at 2 cm H20 and was further enhanced to 139.2 ± 26.9 µm 2/frame at 5 cm H2O (Figure 12F). The number of PTCLs per frame also increased with pressure, from 2.9 ± 0.4 at 0.5 cmH20 to 4.1 ± 0.5 and 5.2 ± 0.6 PTCL/frame at 2 and 5 cmH20 respectively (Figure 12G).

Pressure Dependency of Mouse LMC Diastolic Ca2+ Transients

Representative max projection of GCaMP6f signal over 20 s in an IALVs isolated from Myh11CreERT2-GCaMP6f mice in the presence of the L-type blocker nifedipine (1μM) (A) pressurized to 0.5 cmH2O, 2 cmH2O, 5 cmH2O. The local diastolic Ca2+ transients persist in the presence of nifedipine and increase with increasing pressure as demonstrated in the whole vessel STMs (B). Particle occurrence maps highlight the Ca2+ activity in each LMC as pressure is raised (C). Representative particle analysis plots for particle area (D) and particle counts/frame at each pressure (E). Summary files for particle area (F) and count /frame (G0. * Denotes p<0.05, Mean and SEM shown with n=12 separate IALVs from 8 MYH11-CreERT2-GCaMP6f

Contraction Frequency is Set by the Diastolic Depolarization Rate

Earlier version: To assess how pressure regulates LMC membrane potential we first recorded membrane potential in cells exhibiting action potentials (APs) using a microelectrode filled with biocytin-AF488 to label the cell. In each case (n=3 IALVs) the labeled cell were LMCs wrapping circumferential around the vessel (Figure 13A, B) and as these recordings were made over the course of many minutes the neighboring circumferential LMCs also exhibited fluorescence, albeit weaker in intensity, as expected (Figure 13A). In all the recorded cells exhibiting APs, we noted a diastolic depolarization preceding the sharp upstroke achieved once threshold was met at each pressure (Figure 13C). The AP frequency and rate of the diastolic depolarization increased with pressure (Figure 13D, E). Linear regression of a plot of each AP frequency and diastolic depolarization rate at each pressure demonstrated a tight associated between the two parameters. However, we did not observe a significant effect of pressure on minimum membrane potential (Figure 13G), threshold potential (Figure 13H), the AP upstroke (Figure 13I), AP peak potential (Figure 13J), plateau potential (Figure 13K), or the time spent over threshold (Figure 13L).

Revised version: To assess how pressure regulates LMC membrane potential, we first recorded membrane potential in cells exhibiting action potentials (APs) using a microelectrode filled with biocytin-AF488 to label the impaled cell. In each case (n=3 IALVs) the labeled cell was a LMC wrapping circumferentially around the vessel (Figure 13A, B) and as these recordings were made over the course of many minutes the direct neighboring circumferential LMCs also exhibited fluorescence, albeit weaker in intensity, as expected for cells couple by gap junctions (Figure 13A). In all the recorded cells exhibiting APs, we noted a diastolic depolarization preceding the sharp upstroke achieved once threshold was met at each pressure (Figure 13C). The AP frequency and rate of the diastolic depolarization increased with pressure (Figure 13D, E). Linear regression of a plot of each AP frequency and diastolic depolarization rate at each pressure demonstrated a tight associated between the two parameters. However, we did not observe a significant effect of pressure on minimum membrane potential (Figure 13G), threshold potential (Figure 13H), AP upstroke (Figure 13I), AP peak potential (Figure 13J), plateau potential (Figure 13K), or the time spent over threshold (Figure 13L).

Pressure-Dependent Diastolic Depolarization in LMCs

Intracellular recordings of LMC action potentials (AP) were confirmed by loading (greater than 10minutes) the impaling electrode with 1M KCl 100ug/ml AF488-Biocytin while recording APs followed by imaging on a spinning disk confocal microscope. 3D reconstruction of the z-stack confirmed the circumferential pattern of the impaled LMC that was strongly labeled by AF488-Biocytin (A, B), which also labeled neighboring LMCs, likely through gap junctions as AF488-Biocytin is <1kDa. In a separate set of experiments APs were recorded at 3 different pressures, 0.5 cmH2O, 2 cmH2O, and 5cmH2O. We plotted the representative recordings from 1 cell at each pressure (C). AP frequency was significantly increased with pressure (D) as was the diastolic depolarization rate. Plotting the AP frequency and diastolic depolarization rate from all recordings at each pressure (F) highlights the significant effect diastolic depolarization rate has on the AP frequency. Minimum membrane potential (G), threshold membrane potential of AP initiation (H), upstroke constant (I), peak membrane potential (J), plateau membrane potential (K), and time over threshold (L) are also reported, although not significant.

Discussion

Earlier version: The identification of the cellular origin and signaling mechanisms underlying cLV pacemaking will reveal novel targets for pharmacological intervention in treating lymphedema and the associated lymphatic contractile dysfunction. In this study we tested proposed pacemaker cell types based on 3 parameters: 1) that the pacemaker cells are located along the entire length of the cLV, to accommodate spontaneous contractions and coordinated electrical conduction despite progressive shortening of cLVs; 2) that depolarization of the pacemaker cell can drive a coordinated and propagated contraction of the vessel; and 3) that the presence of Ca2+ transients precedes or coincides with contraction, as commonly observed in other pacemaker cells. We used confocal microscopy and a combination of immunofluorescence and fluorescent reporters under the control of various inducible Cres to identify and target both muscle and non-muscle cells, previously labeled as ICLCs, that coexpress the markers CD34 and PDGFRα. Our cell characterizations were supplemented by scRNAseq analysis of isolated and cleaned mouse IALVs which supported our finding of 3 major cell types including LECs, LMCs, and AdvCs each of which could be further sub-clustered into transcriptionally unique populations. From our initial fluorescence imaging studies, a role for intrinsic pacemaking by LMCs (Van Helden, 1993; von der Weid et al., 2008), or by a novel population of CD34+ lymphatic ICLC (McCloskey et al., 2002; Briggs Boedtkjer et al., 2013), also referred to as telocytes, were further examined and found to co-express Pdgfrα. We utilized PdgfrαCreERTM to further test whether these cells exhibited pacemaker capabilities. However, these Pdgfrα+ cells had minimal Ca2+ activity despite ongoing contractions and optogenetic stimulation of ChR2 in these cells failed to drive a spontaneous contraction. In contrast, photo-stimulation of LMCs expressing ChR2 elicited robust, propagated contractions with similar characteristics and propagation to spontaneous contractions in the same vessels. Furthermore, Ca2+ imaging in LMCs revealed diastolic Ca2+ transients in diastole that increased in frequency and spatial spread as pressure was elevated. We also demonstrated that the primary component of the AP driving the frequency change with pressure is the diastolic depolarization, which we have previously reported to be dependent on Ano1 (Zawieja et al., 2019) and IP3R1 (Zawieja et al., 2023). Notably, we recently reported that diastolic Ca2+ transients are abrogated in IALVs from Myh11CreERT2-Itpr1 inducible knockout mice, supporting a IP3R1-Ano1 axis as the pressure dependent pacemaker mechanism in LMCs. These results, in addition to the recent findings using targeted deletion of Cx45 (Castorena-Gonzalez et al., 2018b) or Cav1.2 (To et al., 2020; Davis et al., 2022) in lymphatic muscle support the model of LMCs as the intrinsic pacemaker as has been previously proposed (Van Helden, 1993; Van Helden et al., 1996; Van Helden and Zhao, 2000).

Revised version: The identification of the cellular origin and signaling mechanisms underlying cLV pacemaking will reveal novel targets for pharmacological intervention in treating lymphedema and the associated lymphatic contractile dysfunction. In this study we tested proposed pacemaker cell types based on 3 parameters: 1) that the pacemaker cells are located along the entire length of the cLV, to accommodate spontaneous contractions and coordinated electrical conduction; 2) that depolarization of the pacemaker cell can drive a coordinated and propagated contraction of the vessel; and 3) that the presence of Ca2+ transients precedes or coincides with contraction, as commonly observed in other pacemaker cells. We used confocal microscopy and a combination of immunofluorescence and fluorescent reporters under the control of various inducible Cres to identify and target both muscle and non-muscle cells, previously labeled as ICLCs, which co-express the markers CD34 and PDGFRα. Our cell characterizations were supplemented by scRNAseq analysis of isolated and cleaned mouse IALVs which supported our finding of three major cell types including LECs, LMCs, and AdvCs each of which could be further sub-clustered into transcriptionally unique populations. From our initial fluorescence imaging studies, a role for intrinsic pacemaking by LMCs (Van Helden, 1993; von der Weid et al., 2008), or by a novel population of CD34+ lymphatic ICLC (McCloskey et al., 2002; Briggs Boedtkjer et al., 2013), also referred to as telocytes, were further examined and found to co-express Pdgfrα. We utilized PdgfrαCreERTM to further test whether these cells exhibited pacemaker capabilities. However, these PDGFRα+ cells had minimal Ca2+ activity despite ongoing contractions and optogenetic stimulation of ChR2 in these cells failed to drive a spontaneous contraction. In contrast, photo-stimulation of LMCs expressing ChR2 elicited robust, propagated contractions with similar characteristics and propagation to spontaneous contractions in the same vessels. Furthermore, Ca2+ imaging in LMCs revealed diastolic Ca2+ transients in diastole that increased in frequency and spatial spread as pressure was elevated. We also demonstrated that the primary component of the AP driving the frequency change with pressure is diastolic depolarization, which we have previously reported to be dependent on ANO1 (Zawieja et al., 2019) and IP3R1 (Zawieja et al., 2023). Notably, we recently reported that diastolic Ca2+ transients are abrogated in IALVs from Myh11CreERT2-Itpr1 inducible knockout mice, supporting a IP3R1-ANO1 axis as the pressure-dependent pacemaker mechanism in LMCs. These results, in addition to the recent findings using targeted deletion of Cx45 (Castorena-Gonzalez et al., 2018b) or Cav1.2 (To et al., 2020; Davis et al., 2022) in lymphatic muscle support the model of LMCs as the intrinsic pacemaker as has been previously proposed (Van Helden, 1993; Van Helden et al., 1996; Van Helden and Zhao, 2000).

Pacemaking in Smooth Muscle

Earlier version: In many smooth muscle organs, regulation of a coordinated contraction is a complex and multicellular phenomenon. Multiple cell types integrate physical and biological information into electrical activity to be transmitted to the force-producing smooth muscle cells, sometimes across great distances relative to cell size, to regulate Ca2+ influx by voltage dependent Ca2+ channels required for contraction. The intestine is one such documented tissue in which cKit+ ICCs and interstitial PDGFRα+ cells form an electrical syncytium to regulate intestinal motility (Sanders et al., 1999; Sanders et al., 2014). The pacemaking function of intestinal ICCs relies heavily on Ano1, a Ca2+ activated Cl- channel, which is required for slow wave activity in the ICCs. Both cKit and Ano1 can be used as a marker for ICCs in the intestine (Hwang et al., 2009; Cobine et al., 2017; Malysz et al., 2017), cKit+ and vimentin+ ICLCs have been observed in sheep lymphatic vessels (McCloskey et al., 2002), yet these cell populations did not form gap junctions with the smooth muscle to form electrical connections (Briggs Boedtkjer et al., 2013) as occurs in the intestinal ICCs. Our cKit staining and CkitCreERT2-ROSA26mTmG reporter studies on mouse IALVs revealed a sparse population of large ovoid cells previously classified as mast cells (Chatterjee and Gashev, 2012; Zawieja et al., 2019). Their identity as mast cells was further supported by a sustained global Ca2+ event after stimulation with the mast cell degranulating agent compound 48-80. However, both vimentin and CD34 showed robust staining throughout the mouse lymphatic vessel wall. Vimentin stained LECs, as well as non-muscle stellate shaped cells, with many co-expressing CD34, and other smaller circular cells some of which were cKit+ as well and some whose morphology was similar to that of the macrophage staining profile of the GFP+ cells in IALVs from MacGreen mice, consistent with previous reports of macrophage staining in cLVs (Bridenbaugh et al., 2013b; Chakraborty et al., 2015; Zawieja et al., 2016). While vimentin+ cells had a perinuclear staining profile, CD34 demarcated the cell membrane and was useful for assessing the morphology of these cells. Of particular interest, the vimentin+CD34+ cells densely populated the length of the mouse IALV, with a majority displaying a flattened stellate morphology characterized by the classic rounded oak leaf appearance, although some displayed fine dendrite extensions. Contrasting with the previous findings in the human thoracic duct (Briggs Boedtkjer et al., 2013), we did not observe a significant population of CD34+ cells with a bipolar morphology oriented axially along the vessel. However, z-stack reconstructions of sections of the mouse IALV that included the secondary valves revealed interstitial CD34+PDGFRα+ cells that resembled the bipolar cells with multiple axon-like extensions throughout the endothelial leaflets similar to interstitial cells that were reported in in collecting vessel valves (Leak and Burke, 1968) and lymphovenous valves (Geng et al., 2016). While these cells have not been frequently described in the peripheral cLV valves, we observed these cells in each of the valve regions imaged in addition to labeling them with other Cre drivers, including NG2Cre-ROSA26mTmG and PdgfrβCreERT2-ROSA26mTmG (data not shown). Whether these cells regulate leaflet extracellular matrix deposition or lymphatic valve integrity is unknown, but a role as a critical pacemaker can be excluded as vessel segments without valves display normal contractile behavior (Van Helden, 1993; Gashev et al., 2002). Instead, the majority of the CD34+PDGFRα+ cells were found in the adventitia with 2-3 layers overtop the LMCs and were consistently observed in high density along the IALV. Some CD34+PDGFRα+ cells or their extensions were present between the lymphatic endothelial and muscle layers as had been previously reported with electron microscopy of human lymphatic vessels (Briggs Boedtkjer et al., 2013). Thus while some of these AdvCs may be contained within the extracellular matrix in the immediate vicinity of the lymphatic vessel, and thus collected during dissection, many others are intimately dispersed within the vessel wall.

Revised version: In many smooth muscle organs, regulation of a coordinated contraction is a complex and multicellular phenomenon. Multiple cell types integrate physical and biological information into electrical activity to be transmitted to the force-producing smooth muscle cells, sometimes across great distances relative to cell size, to regulate Ca2+ influx by voltage-dependent Ca2+ channels required for contraction. The intestine is one such documented tissue in which cKIT+ ICCs and PDGFRα+ interstitial cells form an electrical syncytium to regulate intestinal motility (Sanders et al., 1999; Sanders et al., 2014). The pacemaking function of intestinal ICCs relies heavily on ANO1, a Ca2+ activated Cl- channel, which is required for slow wave activity in the ICCs. Both cKIT and ANO1 can be used as a marker for ICCs in the intestine (Hwang et al., 2009; Cobine et al., 2017; Malysz et al., 2017), cKIT+ and VIMENTIN+ ICLCs have been observed in sheep lymphatic vessels (McCloskey et al., 2002), yet these cell populations did not form gap junctions with the smooth muscle to form electrical connections (Briggs Boedtkjer et al., 2013) as occurs in the intestinal ICCs. Our cKIT staining and cKitCreERT2-ROSA26mTmG reporter studies on mouse IALVs revealed a sparse population of large ovoid cells previously classified as mast cells (Chatterjee and Gashev, 2012; Zawieja et al., 2019). Their identity as mast cells was further supported by a sustained global Ca2+ event after stimulation with the mast cell degranulating agent compound 48-80. However, both VIMENTIN and CD34 showed robust staining throughout the mouse lymphatic vessel wall. LECs stained for VIMENTIN, as did non-muscle stellate shaped cells, with many co-expressing CD34. Other smaller circular cells some of which were also cKIT+ as well and some whose morphology was similar to that of the macrophage staining profile of the GFP+ cells in IALVs from MacGreen mice, were also VIMENTIN+, consistent with previous reports of macrophage staining in cLVs (Bridenbaugh et al., 2013b; Chakraborty et al., 2015; Zawieja et al., 2016). While VIMENTIN+ cells had a perinuclear staining profile, CD34 demarcated the cell membrane and was useful for assessing the morphology of these cells. Of particular interest, the VIMENTIN+CD34+ cells densely populated the length of the mouse IALV, with a majority displaying a flattened stellate morphology characterized by a classic rounded oak-leaf appearance, although some displayed fine dendrite-like extensions. Contrasting with the previous findings in the human thoracic duct (Briggs Boedtkjer et al., 2013), we did not observe a significant population of CD34+ cells with a bipolar morphology oriented axially along the vessel. However, z-stack reconstructions of sections of the mouse IALV that included the secondary valves revealed interstitial CD34+PDGFRα+ cells that resembled those bipolar cells with multiple axon-like extensions throughout the endothelial leaflets; these were similar to interstitial cells that were previously reported in in collecting vessel valves (Leak and Burke, 1968) and lymphovenous valves (Geng et al., 2016). While these cells have not been frequently described in the valves of peripheral cLVs, we observed them in each of the valve regions imaged and, in addition, they were labeled with other Cre drivers, including NG2Cre-ROSA26mTmG and PdgfrβCreERT2-ROSA26mTmG (data not shown). Whether these cells regulate leaflet extracellular matrix deposition or lymphatic valve integrity is unknown, but a possible role as a critical pacemaker can be excluded as vessel segments without valves display normal contractile behavior (Van Helden, 1993; Gashev et al., 2002). Instead, the majority of the CD34+PDGFRα+ cells were found in the adventitia, in 2-3 layers, overtop the LMCs, and they were consistently observed in high density along the IALV. Some CD34+PDGFRα+ cells or their extensions were present between the lymphatic endothelial and muscle layers as had been previously reported with electron microscopy of human lymphatic vessels (Briggs Boedtkjer et al., 2013). Thus, while some of these AdvCs may be contained within the extracellular matrix that retracts onto the vessel during microdissection, many others are intimately dispersed within the vessel wall.

PDGFRα+CD34+ Cells are Not Involved in cLV Pacemaking Under Physiological Conditions

Earlier version: Co-expression of CD34 and PDGFRα has recently been ascribed as a delineating feature of telocytes, although PDGFRα routinely labels fibroblasts and specific interstitial cells in the GI tract involved in purinergic neurotransmission (Kurahashi et al., 2011; Kurahashi et al., 2013; Clayton et al., 2022), and CD34 expression has been observed in multipotent cell populations of various origins (Sidney et al., 2014). In the intestinal villi, PDGFRα+ fibroblasts appear to be the progenitors of the smooth muscle fibers associated with the lacteal, the lymphatic capillary in the villi (Sanketi et al., 2024). Of course, mesenchymal stromal cells (Andrzejewska et al., 2019) and fibroblasts (Muhl et al., 2020; Buechler et al., 2021; Forte et al., 2022), and it remains controversial to what extent telocytes are distinct from or are components/subtypes of either cell type (Clayton et al., 2022), are not monolithic in their expression patterns displaying both organ directed transcriptional patterns as well as intra-organ heterogeneity (Lendahl et al., 2022) as readily demonstrated by recent single cell RNA sequencing studies that provided immense detail about the subtypes and activation spectrum within these cells and their plasticity (Luo et al., 2022a). We were able to distinguish up to 10 subclusters of AdvCs, the majority of which expressed or co-expressed CD34 and PDGFRα. These cells were consistently negative for smooth muscle markers such as Des, Cnn1, Acta2, Myh11 or the pericyte marker Mcam. However, PDGFRβ expression was noted in our scRNAseq and sorted PdgfrαCreERTM-ROSA26mTmG cells, in addition to the staining of LMCs. PDGFRβ protein expression was confirmed with variable immunofluorescence staining amongst the PDGFRα stained cells as well as LMCs. The PdgfrβCreERT2ROSA26mTmG mice had only modest recombination in both the LMC and PDGFRα+ cell population, but potentially highlighted a myofibroblast-like cell subpopulation, cells that might lie on the spectrum of differentiation from lymphatic muscle and PDGFRα+ cells, or perhaps a cell with pacemaker activity as PDGFRβ is widely used as a pericyte marker and some pericytes display pacemaker activity (Hashitani et al., 2015). Adding to this intrigue, the PdgfrαCreERTM sorted cells expressed transcripts for Cacna1c, the voltage-gated L-type Ca2+ channel critical for lymphatic contractions (Zawieja et al., 2018a; To et al., 2020); Ano1, the ion channel underlying pressure dependent chronotropy (Mohanakumar et al., 2018; Zawieja et al., 2019); and Cx45, the primary connexin mediating electrical conduction in mouse lymphatic collecting vessels (Castorena-Gonzalez et al., 2018b; Hald et al., 2018). Expression of these genes in certain sub-populations of the AdvCs was also apparent in our scRNAseq analysis. Thus, the presence of those gene transcripts does not appear to be due to muscle cell contamination or incidental recombination in LMCs as we did not detect LMC markers in the RT-PCR profiling of the sorted PDGFRα+ cells nor were GFP-expressing cells with an LMC morphology observed in imaging of PdgfrαCreERTM-ROSA26mTmG vessels. Critically, however, deletion of Cav1.2, Cx45, or Ano1 through PdgfrαCreERTM-mediated recombination neither recapitulated the previous phenotypes achieved with Myh11CreERT2 (Castorena-Gonzalez et al., 2018b; Zawieja et al., 2019; To et al., 2020; Davis et al., 2022) nor significantly affected pacemaking in mouse popliteal cLVs. This is in stark contrast to the complete lack of contractions observed in Myh11CreERT2-Cav1.2 fl/fl vessels (To et al., 2020) or the vessels from vascular muscle specific Itga8CreERT2-Cav1.2fl/fl mice (Davis et al., 2022; Warthi et al., 2022), and the significant loss in pressure-induced chronotropic modulation of pacemaker function in IALVs with Myh11CreERT2-mediated deletion of Ano1 that we have previously reported (Zawieja et al., 2019). While a sub-population of CD34+PDGFRα+ cells may share expression of critical pacemaker genes identified in the LMCs, they do not appear to be involved in cLV pacemaking or contractile function under physiological states. Instead, CD34+PDGFRα+ cells co-stained significantly with Sca1+, suggesting they may be primed to act as resident multipotent cells (Song et al., 2020; Kimura et al., 2021). To this point, the PdgfrαCreERTM FACS purified cells also expressed markers associated with “stemness” such as CD34, Klf4, Gli1, CD29, CD105, CD44, and Vimentin, in addition to Sca1, and it is likely that the PdgfrαCreERTM population includes various distinct subpopulations (Jolly et al., 2022) expressing these markers. These cells may play a role in rebuilding the lymphatic collecting vessel vasculature following collecting vessel damage or lymph node resection and further studies are required to assess their functional multipotency.

Revised version: Co-expression of CD34 and PDGFRα has recently been ascribed as a delineating feature of telocytes, although PDGFRα routinely labels fibroblasts and specific interstitial cells in the GI tract involved in purinergic neurotransmission in the GI tract (Kurahashi et al., 2011; Kurahashi et al., 2013; Clayton et al., 2022). CD34 expression is also ascribed to some multipotent cell populations of various origins (Sidney et al., 2014). For example, PDGFRα+ fibroblasts appear to be progenitors of the smooth muscle fibers associated with the lacteal, the lymphatic capillary in the villus (Sanketi et al., 2024). It remains controversial to what extent telocytes are distinct from or are components/subtypes of either cell type and morphological discrimination between the populations typically requires electron microscopy imaging (Clayton et al., 2022). Mesenchymal stromal cells (Andrzejewska et al., 2019) and fibroblasts (Muhl et al., 2020; Buechler et al., 2021; Forte et al., 2022) are not monolithic in their expression patterns displaying both organ directed transcriptional patterns as well as intra-organ heterogeneity (Lendahl et al., 2022) as readily demonstrated by recent single cell RNA sequencing studies that provided immense detail about the subtypes and activation spectrum within these cells and their plasticity (Luo et al., 2022a). We were able to distinguish up to 10 subclusters of AdvCs, the majority of which expressed or co-expressed CD34 and PDGFRα. These cells were consistently negative for smooth muscle markers such as Des, Cnn1, Acta2, Myh11 or the pericyte marker Mcam. However, PDGFRβ expression was noted in our scRNAseq analysis and in our RT-PCR of sorted PdgfrαCreERTM-ROSA26mTmG cells. PDGFRβ protein expression was confirmed with variable immunofluorescence staining amongst the PDGFRα stained cells as well as LMCs. The PdgfrβCreERT2ROSA26mTmG mice had only modest recombination in both the LMC and PDGFRα+ cell populations, but potentially highlighted a myofibroblast-like cell subpopulation, cells that might lie on the spectrum of differentiation from lymphatic muscle and PDGFRα+ cells, or perhaps a cell with pacemaker activity as PDGFRβ is widely used as a pericyte marker and some pericytes display pacemaker activity (Hashitani et al., 2015). Adding to this intrigue, the PdgfrαCreERTM sorted cells expressed transcripts for Cacna1c, the voltage-gated L-type Ca2+ channel critical for lymphatic contractions (Zawieja et al., 2018a; To et al., 2020); Ano1, the ion channel underlying pressure-dependent chronotropy (Mohanakumar et al., 2018; Zawieja et al., 2019); and Cx45, the primary connexin mediating electrical conduction in mouse lymphatic collecting vessels (Castorena-Gonzalez et al., 2018b; Hald et al., 2018). Expression of these genes in certain sub-populations of the AdvCs was also apparent in our scRNAseq analysis. Thus, the presence of those gene transcripts does not appear to be due to muscle cell contamination or incidental recombination in LMCs as we did not detect LMC markers in the RT-PCR profiling of the sorted PDGFRα+ cells nor were GFP-expressing cells with an LMC morphology observed in imaging of PdgfrαCreERTM-ROSA26mTmG vessels. Critically, however, deletion of Cav1.2, Cx45, or Ano1 through PdgfrαCreERTM-mediated recombination neither recapitulated the previous phenotypes achieved with Myh11CreERT2(Castorena-Gonzalez et al., 2018b; Zawieja et al., 2019; To et al., 2020; Davis et al., 2022) nor significantly affected pacemaking in mouse popliteal cLVs. This finding is in stark contrast to the complete lack of contractions observed in Myh11CreERT2-Cav1.2 fl/fl vessels (To et al., 2020; Davis et al., 2023b) or the vessels from vascular muscle specific Itga8CreERT2-Cav1.2fl/fl mice (Davis et al., 2022; Warthi et al., 2022), and the significant loss in pressure-induced chronotropic modulation of pacemaker function in IALVs with Myh11CreERT2-mediated deletion of Ano1 that we have previously reported (Zawieja et al., 2019). While a sub-population of CD34+PDGFRα+ cells may share expression of critical pacemaker genes identified in the LMCs, they do not appear to be involved in cLV pacemaking or contractile function under physiological states. Instead, CD34+PDGFRα+ cells co-stained significantly with Sca1+, suggesting they may be primed to act as resident multipotent cells (Song et al., 2020; Kimura et al., 2021). To this point, the PdgfrαCreERTM FACS purified cells also expressed markers associated with “stemness” such as CD34, Klf4, Gli1, CD29, CD105, CD44, and Vimentin, in addition to Sca1, and it is likely that the PdgfrαCreERTM population includes various distinct subpopulations (Jolly et al., 2022) expressing these markers. These cells may play a role in rebuilding the lymphatic collecting vessel vasculature following collecting vessel damage or lymph node resection and further studies are required to assess their functional multipotency.

SR Ca2+ Cycling in Pacemaking

Earlier version: The use of the mouse model, in addition to the simplicity of the vessel architecture, provided the use of genetic tools that previously had been instrumental in identifying the cKit+ ICC as the pacemaker cells of the GI tract (Ward et al., 1994; Huizinga et al., 1995; Torihashi et al., 1995). Through the use of the respective PdgfrαCreERTM and Myh11CreERT2 models, we were able to specifically image Ca2+ in each cell type in pressurized, contracting vessels. Pacemaking initiating cells have an inherently unstable membrane potential, oftentimes utilizing the oscillatory nature of Ca2+ release from the sarcoendoplasmic reticulum coupled to Ca2+ sensitive electrogenic exchangers and ion channels to drive depolarization (Van Helden, 1993; Hashitani et al., 2015; Baker et al., 2021b; Sanders et al., 2022). One such example is the pacemaker ICC in the gastric corpus which exhibits abundant Ca2+ transients that couple to Ano1-mediated chloride currents in both the intervening period between slow waves as well as the plateau phase of the slow wave (Baker et al., 2021a), although such activity is not characteristic of all pacemaker ICC types. The identification of a Ca2+ activated chloride current in LMCs (Van Helden, 1993; Toland et al., 2000) and its correspondence with subcellular Ca2+ transients (Van Helden, 1993; Ferrusi et al., 2004; von der Weid et al., 2008) led Van Helden to postulate that LMCs had an intrinsic pacemaking capability (Van Helden, 1993; Van Helden et al., 1996). We have previously reported that mouse LMCs in pressurized vessels routinely display subcellular Ca2+ release events that reflect the kinetics and characteristics of Ca2+ puffs and waves in addition to the coordinated global Ca2+ flash associated with Ca2+ influx during an AP (Castorena-Gonzalez et al., 2018b; Zawieja et al., 2018a; Zawieja et al., 2019). Here we confirmed the consistent presence of subcellular Ca2+ transients only in LMCs with GCaMP6f driven by Myh11CreERT2 but not in the cells with GCaMP6f driven by PdgfrαCreERTM. Critically, we also demonstrated that the Ca2+ transients increased in both frequency and spatial spread as pressure was elevated in the vessel, as would be expected to account for the pressure dependent chronotropy observed in lymphatic collecting vessels. This underscores the recent finding that the genetic deletion of Ano1 in the LMCs dramatically reduced contraction frequency and abolished pressure-dependent chronotropy in those vessels (Zawieja et al., 2019). This phenotype was largely replicated with a similar reduction in frequency and loss of pressure dependent chronotropy in our recent study utilizing Myh11CreERT2 to drive deletion of IP3R1 from LMCs (Zawieja et al., 2023) in which these diastolic Ca2+ transients were absent. This fits with the central role of IP3R and subcellular Ca2+ release as critical components of intrinsic LMC pacemaking (Van Helden et al., 1996; von der Weid et al., 2008). In addition to the transcriptional heterogeneity identified by scRNASeq, we also noted heterogeneity in the propensity of LMCs to display diastolic Ca2+ transients under control conditions or the sustained Ca2+ oscillations that occur in the presence of nifedipine. We did not detect and significant difference in the expression of Itpr1, the gene encoding the IP3R1, across our LMCs subclusters. However, the LMC cluster “0” had significantly increased expression of Itprid2 which encodes the KRas-induced actin-interacting protein (KRAP). KRAP has recently been implicated in IP3R1 immobilization and licensing and was required for IP3R1-mediated Ca2+ puffs (Thillaiappan et al., 2021; Atakpa-Adaji et al., 2024). Whether the higher expression of KRAP results in a greater probability of these LMCs to display IP3R1-dependent Ca2+ oscillations in LMCs requires further investigation.

Revised version: The use of the mouse IALV model, in addition to the simplicity of the vessel architecture, provided the use of genetic tools that previously had been instrumental in identifying the cKIT+ ICC as the pacemaker cells of the GI tract (Ward et al., 1994; Huizinga et al., 1995; Torihashi et al., 1995). Through the use of the respective PdgfrαCreERTM and Myh11CreERT2 drivers, we were able to specifically image Ca2+ in each respective cell type in pressurized, contracting vessels. Pacemaking initiating cells have an inherently unstable membrane potential, oftentimes utilizing the oscillatory nature of Ca2+ release from the sarcoendoplasmic reticulum coupled to Ca2+ sensitive electrogenic exchangers or ion channels to drive depolarization (Van Helden, 1993; Hashitani et al., 2015; Baker et al., 2021b; Sanders et al., 2022). One such example is the pacemaker ICC in the gastric corpus which exhibits abundant Ca2+ transients that couple to ANO1-mediated chloride currents in both the intervening period between slow waves as well as the plateau phase of the slow wave (Baker et al., 2021a), however such activity is not characteristic of all pacemaker ICC types. The identification of a Ca2+-activated chloride current in LMCs (Van Helden, 1993; Toland et al., 2000) and its correspondence with subcellular Ca2+ transients (Van Helden, 1993; Ferrusi et al., 2004; von der Weid et al., 2008) led Van Helden to postulate that LMCs had intrinsic pacemaking capability (Van Helden, 1993; Van Helden et al., 1996). We have previously reported that mouse LMCs in pressurized vessels routinely display subcellular Ca2+ release events that reflect the kinetics and characteristics of Ca2+ puffs and waves in addition to the coordinated global Ca2+ flash associated with Ca2+ influx during an AP (Castorena-Gonzalez et al., 2018b; Zawieja et al., 2018a; Zawieja et al., 2019). Here we confirmed the consistent presence of subcellular Ca2+ transients only in LMCs with GCaMP6f driven by Myh11CreERT2 but not in the cells with GCaMP6f driven by PdgfrαCreERTM. Critically, we also demonstrated that the Ca2+ transients increased in both frequency and spatial spread as pressure was elevated in the vessel, as would be expected to account for the pressure-dependent chronotropy observed in lymphatic collecting vessels. This underscores the finding that genetic deletion of Ano1 in the LMCs dramatically reduced contraction frequency and abolished pressure-dependent chronotropy in those vessels (Zawieja et al., 2019). This phenotype was largely replicated with a similar reduction in frequency and loss of pressure-dependent chronotropy in our recent study utilizing Myh11CreERT2 to drive deletion of IP3R1 from LMCs (Zawieja et al., 2023) in which these diastolic Ca2+ transients were absent. This fits with the central role of IP3R and subcellular Ca2+ release as critical components of intrinsic LMC pacemaking (Van Helden et al., 1996; von der Weid et al., 2008). In addition to the transcriptional heterogeneity identified by scRNASeq, we also noted heterogeneity in the propensity of LMCs to display diastolic Ca2+ transients under control conditions or show the sustained Ca2+ oscillations that occur in the presence of nifedipine. We did not detect any significant difference in the expression of Itpr1, the gene encoding the IP3R1, across our LMCs subclusters. However, when using less stringent conditions we identified that the LMC cluster “0” had significantly increased expression of Itprid2 (Log2FC of 0.26) which encodes the KRas-induced actin-interacting protein (KRAP). KRAP has recently been implicated in IP3R1 immobilization and licensing and was required for IP3R1-mediated Ca2+ puffs (Thillaiappan et al., 2021; Atakpa-Adaji et al., 2024). Whether the higher expression of KRAP results in a greater probability of these LMCs to display IP3R1-dependent Ca2+ oscillations in LMCs requires further investigation. Of note, LMCs also express the components for store operated calcium entry including Stim1, Stim2, Orai1, Orai3, Saraf, and Stimate, which may be involved in maintaining IP3R1-dependent SR Ca2+ release oscillations.

Earlier version: The membrane potential recordings we made in this work suggest that the regulation of pressure dependent chronotropy is through modulation of the diastolic depolarization rate in LMCs, as previously demonstrated in rat mesenteric lymphatic vessels (Zawieja et al., 2018b). The appearance of the diastolic depolarization may depend on the method of stretch employed as it is not always observed in preparations using a wire myograph (von der Weid et al., 2014). Notably, in this study PdgfrαCreERTM mediated deletion of Ano1 had no effect on contractile parameters. The lack of Ca2+ transients in PDGFRα+ cells across any stage of the lymphatic contraction cycle diminishes any expected role for this cell type to perform as the pacemaker for the mouse IALV. In the pacemaker ICCs, slow wave activity is dependent on calcium transient clusters Instead pressure-dependent Ca2+ mobilization from the SR, through IP3R1 (Zawieja et al., 2023), appears to set the basis for LMC pacemaking as previously proposed (Van Helden, 1991; von der Weid et al., 2008). However, the mechanisms driving IP3R1 activation and Ca2+ oscillations remain to be fully addressed.

Revised version: The membrane potential recordings we made in this study suggest that the regulation of pressure-dependent chronotropy is through modulation of the diastolic depolarization rate in LMCs, as previously demonstrated in rat mesenteric lymphatic vessels (Zawieja et al., 2018b). The appearance of the diastolic depolarization may depend on the method of vessel stretch employed as it is not always observed in wire myograph preparations (von der Weid et al., 2014). Notably, in this study PdgfrαCreERTM mediated deletion of Ano1 had no effect on contractile parameters. The lack of Ca2+ transients in PDGFRα+ cells across any stage of the lymphatic contraction cycle also diminishes any expected role for this cell type to perform as the pacemaker for the mouse IALV. We recently showed that pressure-dependent Ca2+ mobilization from the SR, through IP3R1 (Zawieja et al., 2023), sets the basis for LMC pacemaking as previously proposed (Van Helden, 1991; von der Weid et al., 2008). However, the mechanisms driving IP3R1 activation and Ca2+ oscillations remain to be fully addressed.

Earlier version: A pacemaker cell would be expected to be electrically coupled to the LMC layer to permit the nearly synchronous conduction velocity of the contraction wave (Zawieja et al., 1993; Castorena-Gonzalez et al., 2018b; Hald et al., 2018) and to transmit depolarization into coupled LMCs to activate the voltage dependent Ca2+ channels that are responsible for lymphatic muscle APs. Connexins are the molecular constituents of gap junctions and, as stated above, we detected Cx45 expression in the PdgfrαCreERTM sorted cells. However, we did not detect any impairment in pacemaking, nor were contraction conduction speed deficits or multiple pacemakers noted in the PdgfrαCreERTM -Cx45fl/fl popliteal cLVs, in contrast to the development of multiple pacemaker sites and the lack of entrainment that characterize cLVs from Myh11CreERT2-Cx45fl/fl mice (Castorena-Gonzalez et al., 2018b). Admittedly, we did not perform an exhaustive assessment of the connexin expression profile of the CD34+PDGFRα+ cells, and Cx45 may not be the dominant connexin expressed in the CD34+PDGFRα+ cells, or heterotypic connexons could exist (Koval et al., 2014). However, electron microscopy studies of the putative ICLC in the human thoracic duct did not detect any gap junctions, although peg-and-socket connections were observed (Briggs Boedtkjer et al., 2013). We utilized optogenetics to directly depolarize the specific cell populations in both the PdgfrαCreERTM and Myh11CreERT2 mouse models in an attempt to drive contractions. Local photo-stimulation of the PDGFRα cells failed to initiate contraction while the stimulation of Myh11CreERT2 recombined cells resulted in contractions that were indistinguishable from the spontaneously occurring contractions. These results give functional credence to the lack of hetero-cellular coupling via gap junctions that was previously reported (Briggs Boedtkjer et al., 2013). Just as critically, they also highlight the regenerative nature of the lymphatic muscle AP as local depolarization was sufficient to drive a coordinated contraction along the vessel and that a single or few LMCs reaching threshold for AP initiation are sufficient to drive the conducted activity observed at the tissue level.

Revised version: A pacemaker cell would be expected to be electrically coupled to the LMC layer to permit the nearly synchronous conduction velocity of the contraction wave (Zawieja et al., 1993; Castorena-Gonzalez et al., 2018b; Hald et al., 2018) and to transmit depolarization into coupled LMCs to activate the voltage dependent Ca2+ channels that are responsible for lymphatic muscle APs. Connexins are the molecular constituents of gap junctions and, as stated above, we detected Cx45 expression in PdgfrαCreERTMsorted cells. However, we did not detect any impairment in pacemaking, nor were contraction conduction speed deficits or multiple pacemakers noted in the PdgfrαCreERTM-Cx45fl/fl popliteal cLVs, in contrast to the development of multiple pacemaker sites and the lack of entrainment that characterize cLVs from Myh11CreERT2-Cx45fl/fl mice (Castorena-Gonzalez et al., 2018b). Admittedly, we did not perform an exhaustive assessment of the connexin expression profile of the CD34+PDGFRα+ cells, and Cx45 may not be the dominant connexin expressed in the CD34+PDGFRα+ cells, or heterotypic connexons could exist (Koval et al., 2014). However, electron microscopy studies of the putative ICLC in the human thoracic duct did not detect any gap junctions, although peg-and-socket connections were observed (Briggs Boedtkjer et al., 2013). We utilized optogenetics to directly depolarize the specific cell populations in both the PdgfrαCreERTM and Myh11CreERT2 mouse models in an attempt to drive contractions. Local photo-stimulation of the PDGFRα cells failed to initiate contraction while the stimulation of Myh11CreERT2 recombined cells resulted in contractions that were indistinguishable from the spontaneously occurring contractions. These results give functional credence to the lack of hetero-cellular coupling via gap junctions that was previously reported (Briggs Boedtkjer et al., 2013). Just as critically, our results also highlight the regenerative nature of the lymphatic muscle AP. Local, optogenetic-initiated depolarization of either a single or a few LMCs to threshold was sufficient to drive a coordinated contraction along the vessel conducted activity at the tissue level.

Conclusions

Earlier version: Our present findings lend further support to the hypothesis that the LMCs are intrinsic pacemakers (van Helden et al., 2006; Mitsui and Hashitani, 2020) and mouse cLVs do not require an ICC-like cell network to drive propagated contractions. These findings also underscore the significance of lymphatic muscle Ca2+ handling as the driver of lymphatic pacemaking, which can be compromised in disease states leading to impaired lymphatic contractile activity (Stolarz et al., 2019; Lee et al., 2020; Van et al., 2021). Further studies delineating the specific SR Ca2+ release and influx pathways, and the contributions of Ca2+ sensitive ion channels are required to develop sophisticated in silico models and identify potential therapeutic targets to rescue lymphatic pacemaking in lymphedema patients (Olszewski, 2002, 2008).

Revised version: Our present findings lend further support to the hypothesis that the LMCs are intrinsic pacemakers (van Helden et al., 2006; Mitsui and Hashitani, 2020) and that mouse cLVs do not require an ICC-like cell network to drive propagated contractions. These findings also underscore the significance of lymphatic muscle Ca2+ handling as the driver of lymphatic pacemaking, which can be compromised in disease states leading to impaired lymphatic contractile activity (Stolarz et al., 2019; Lee et al., 2020; Van et al., 2021). Further studies delineating the specific SR Ca2+ release and influx pathways, and the contributions of Ca2+ sensitive ion channels are required to develop sophisticated in silico models and identify potential therapeutic targets to rescue lymphatic pacemaking in lymphedema patients (Olszewski, 2002, 2008).

Limitations

Earlier version: One fundamental assumption underlying our conclusions is that there is a conserved pacemaking pathway and cell type regulating lymphatic collecting vessel contractions across species, specifically pertaining to the capability of lymphatic muscle to maintain pacemaking and coordination despite changes in tissue complexity and cLV wall thickness. It is worth noting that lymphatic collecting vessels in mice have similar pressure-dependent chronotropy and contraction conduction velocity as recorded in rats and human vessels (Castorena-Gonzalez et al., 2018b). These similarities exist despite the fact that mouse lymphatic collecting vessels are typically encircled by a single layer of lymphatic muscle while larger species may have multiple layers of LMCs in the wall. It is possible that vessels with multiple layers of LMCs need a network of ICLC to coordinate their activity. The simplicity in the makeup of the mouse cLV and the use of cell targeting Cre models provide great control over experimental variables, but other cell types may provide coordination of LMC pacemaking in other species where the lymphatic cLV walls are larger and thicker and contain multiple muscle cell layers. Our scRNAseq analysis also is likely biased using ROSAmTmG mice with FACS purification to remove debris and concentrate specific cell types from these pooled small vessels. Larger and more complex cells, with attributes that can be ascribed to ICCs, are more likely to be lost in this methodology (e.g.., depending on the FACS gating parameters) and this procedure can also elicit a stress response in the transcriptome of the analyzed cells. However, we also did not observe long and complex cells, aside from the circumferential LMCs, in our immunofluorescence and recombination reporter imaging experiments. Immediate and early gene expression motifs driven by a stress response may be a component of the differences in sub-clusters that were identified. Future scRNAseq or snRNAseq studies with deeper sequencing will be required to ensure the full transcriptomic heterogeneity is accounted under different cellular stress conditions.

Revised version: One fundamental assumption underlying our conclusions is that there is a conserved pacemaking pathway and cell type regulating lymphatic collecting vessel contractions across species, specifically pertaining to the capability of lymphatic muscle to maintain pacemaking and coordination despite changes in tissue complexity and cLV wall thickness. It is worth noting that lymphatic collecting vessels in mice have similar pressure-dependent chronotropy and contraction conduction velocity as recorded in rats and human vessels (Castorena-Gonzalez et al., 2018b). These similarities exist despite the fact that mouse lymphatic collecting vessels are typically encircled by a single layer of lymphatic muscle while larger species may have multiple layers of LMCs in the wall. It is possible that vessels with multiple layers of LMCs need a network of ICLC to coordinate their activity. The simplicity in the makeup of the mouse cLV and the use of cell targeting Cre models provides great control over experimental variables, but other cell types may be required for coordination of LMC pacemaking in other species where the lymphatic cLV walls are larger and thicker and contain multiple muscle cell layers. Our scRNAseq analysis also is likely biased using ROSAmTmG mice with FACS purification to remove debris and concentrate specific cell types from these pooled small vessels. Larger and more complex cells, with attributes that can be ascribed to ICCs, are more likely to be lost in this methodology (e.g.., depending on the FACS gating parameters) and this procedure can also elicit a stress response in the transcriptome of the analyzed cells. However, we also did not observe long and complex cells, aside from the circumferential LMCs, in our immunofluorescence and recombination reporter imaging experiments. Immediate and early gene expression motifs driven by a stress response may be a component of the differences in sub-clusters that were identified. Future scRNAseq or snRNAseq studies with deeper sequencing will be required to ensure that the full transcriptomic heterogeneity is accounted under different cellular stress conditions.

Earlier version: Our data demonstrate that limited staining of a few cell markers alone is insufficient to identify discrete cell populations in the mouse cLVs. Additionally, mRNA expression does not equal protein translation nor guarantee specific function as we did not readily detect endothelial CD34 with immunofluorescence despite detecting transcript; additionally, PdgfrαCreERTM mediated deletion of Ano1, Cx45, or Cav1.2 had no effect on cLV pacemaking. Hence, further experimentation is also required to fully characterize expression of multipotent cell markers and function of CD34+PDGFRα+Sca1+ cells invested within the mouse cLVs, although this was beyond the scope of this study assessing pacemaker identity. Tangentially, another limitation of our approach pertains to the specificity and recombination efficiency of inducible Cre recombinase models, which can be a notable confounding variable (Chakraborty et al., 2019). We observed that our inducible Cre models led to a degree of nonspecific recombination within the mouse cLV, with the GCaMP6f and ChR2 particularly susceptible to recombination compared to the ROSA26mT/mG reporter. Recombination in multiple cell types was expected with the constitutive Cre models we employed (Ng2Cre and PdgfrαCre), as vascular and lymphatic muscle precursor cells can transiently express Nestin, Pdgfrα, and Ng2 (Hill et al., 2015; Castorena-Gonzalez et al., 2018b; Kenney et al., 2020). We also observed that PdgfrβCreERT2 drove recombination in a sub population of LMCs and PDGFRα+ cells. These appeared to be two distinct populations that only share expression for PDGFRβ based on our scRNAseq dataset but may exist along a continuum of differentiation. PDGFB-PDGFRβ signaling is critical for normal mural cell recruitment to both the blood and lymphatic vasculature (Gaengel et al., 2009; Wang et al., 2017) and proliferating vascular smooth muscle cells and pericytes have both been documented to express PDGFRβ (Andrae et al., 2008; Pitulescu and Adams, 2014). Ideally, novel Cre or combinatorial Cre models that specifically target LMCs or sub populations of LMCs may be developed to further tease out the functional role of these cells.

Revised version: Our data demonstrate that limited staining of a few cell markers alone is insufficient to identify discrete cell populations in mouse cLVs. Additionally, mRNA expression does not equal protein translation nor guarantee specific function as we did not readily detect endothelial CD34 with immunofluorescence despite detecting transcript; additionally, PdgfrαCreERTM-mediated deletion of Ano1, Cx45, or Cav1.2 had no effect on cLV pacemaking. Further experimentation is also required to fully characterize expression of multipotent cell markers and function of CD34+PDGFRα+Sca1+ cells invested within the mouse cLVs, although doing so was beyond the scope of this study assessing pacemaker identity. Tangentially, another limitation of our approach pertains to the specificity and recombination efficiency of inducible Cre recombinase models, which can be a notable confounding variable (Chakraborty et al., 2019). We observed that our inducible Cre models led to a degree of nonspecific recombination within the mouse cLV, with GCaMP6f and ChR2 particularly susceptible to recombination compared to the ROSA26mT/mG reporter. Recombination in multiple cell types was expected with the constitutive Cre models we employed (Ng2Cre and PdgfrαCre), as vascular and lymphatic muscle precursor cells can transiently express Nestin, Pdgfrα, and Ng2 (Hill et al., 2015; Castorena-Gonzalez et al., 2018b; Kenney et al., 2020). We also observed that PdgfrβCreERT2 drove recombination in a sub population of LMCs and PDGFRα+ cells. These appeared to be two distinct populations that only share expression for PDGFRβ based on our scRNAseq dataset, but which may exist along a continuum of differentiation. PDGFB-PDGFRβ signaling is critical for normal mural cell recruitment to both the blood and lymphatic vasculature (Gaengel et al., 2009; Wang et al., 2017) and proliferating vascular smooth muscle cells and pericytes have both been documented to express PDGFRβ (Andrae et al., 2008; Pitulescu and Adams, 2014). Ideally, novel Cre or combinatorial Cre models that specifically target LMCs or sub populations of LMCs may be developed to further tease out the functional roles of these cells.

Materials and Methods

Mice

Earlier version: Wild-type (WT) male mice (25-35 g) on the C57BL/6J background, ROSA26mT/mG reporter (Muzumdar et al., 2007) (Strain#007676), transgenic PdgfrαCre (Strain#013148), CSFR1-EGFP (MacGreen) (Sasmono et al., 2003)(Strain#018549), genetically encoded Ca2+ sensor GCaMP6f (Chen et al., 2013) (Strain#028865), transgenic PdgfrαCreERTM (Kang et al., 2010) (Strain#018280), NG2-Cre (Strain #:008533)(Zhu et al., 2008), and ChR2 /tdTomato fusion mice (Madisen et al., 2012) (Strain#012567) were purchased from The Jackson Laboratory (Bar Harbor, MA, USA). PdgfrβCreERT2 (Gerl et al., 2015) mice were a gift from Ralf Adams (Mac Planck Institute) and kindly provided by Lorin Olson (Oklahoma Medical Research Foundation) and are currently available at Jax (Strain#029684). The Myh11CreERT2 mice (Wirth et al., 2008) were a gift from Stefan Offermanns, Max-Planck-Intstitut fur Herz-und Lungendforschung, Bad Nauheim, Germany, and are currently available at Jax (Strain #019079, Y-Linked). c-KitCreERT2 mice (Heger et al., 2014) were a gift from Dieter Saur (Technical University of Munich). Prox1-eGFP mice (Choi et al., 2011) were a gift from Young-Kwon Hong (University of Southern California. For genotyping, we isolated genomic DNA from mouse tail clips using the HotSHOT method (Truett et al., 2000). Specific mouse genotypes were confirmed via PCR using 2x PCR Super Master Polymerase Mix (Catalog # B46019, Bimake, Houston, TX) performed as specified by the provider. Mice used for this study were 3-8 months of age. All animal protocols were approved by the University of Missouri Animal Care and Use Committee and conformed to the US Public Health Service policy for the humane care and use of laboratory animals (PHS Policy, 1996).

Revised version: Wild-type (WT) male mice (25-35 g) on the C57BL/6J background, ROSA26mT/mG reporter (Muzumdar et al., 2007) (Strain#007676), transgenic PdgfrαCre (Strain#013148), CSFR1-EGFP (MacGreen) (Sasmono et al., 2003) (Strain#018549), genetically encoded Ca2+ sensor GCaMP6f (Chen et al., 2013) (Strain#028865), transgenic PdgfrαCreERTM (Kang et al., 2010) (Strain#018280), NG2-Cre (Strain #:008533) (Zhu et al., 2008), and ChR2/tdTomato fusion mice (Madisen et al., 2012) (Strain#012567) were purchased from The Jackson Laboratory (Bar Harbor, MA, USA). PdgfrβCreERT2 (Gerl et al., 2015) mice were a gift from Ralf Adams (Mac Planck Institute) and kindly provided by Lorin Olson (Oklahoma Medical Research Foundation) and are currently available at Jax (Strain#029684). The Myh11CreERT2 mice (Wirth et al., 2008) were a gift from Stefan Offermanns, Max-Planck-Intstitut fur Herz-und Lungendforschung, Bad Nauheim, Germany, and are currently available at Jax (Strain #019079, Y-Linked). c-KitCreERT2mice (Heger et al., 2014) were a gift from Dieter Saur (Technical University of Munich). Prox1-eGFP mice (Choi et al., 2011) were a gift from Young-Kwon Hong (University of Southern California). For genotyping, we isolated genomic DNA from mouse tail clips using the HotSHOT method (Truett et al., 2000). Specific mouse genotypes were confirmed via PCR using 2x PCR Super Master Polymerase Mix (Catalog # B46019, Bimake, Houston, TX) performed as specified by the provider. Mice used for this study were 3-12 months of age. All animal protocols were approved by the University of Missouri Animal Care and Use Committee and conformed to the US Public Health Service policy for the humane care and use of laboratory animals (PHS Policy, 1996).

iCre Tamoxifen Induction

Earlier version: Mice harboring PdgfrαCreERTM, PdgfrβCreERT2, Myh11CreERT2, and cKitCreERT2 were crossed with ROSA26mT/mG mice to generate PdgfrαCreERTM-ROSA26mT/mG, PdgfrβCreERT2-ROSA26mT/mG, Myh11CreERT2-ROSA26mT/mG, and cKitCreERT2-ROSA26mT/mG mice, respectively. The resulting iCre-ROSA26mT/mG mice were induced with tamoxifen 2-4 weeks after weaning. Tamoxifen induction was performed via consecutive 100 μL i.p. injections of tamoxifen ranging from 1 to 5 days at concentrations ranging from 0.2 -10 mg/mL in safflower oil, using a titrated induction protocol to determine the extent of recombination in specific cell populations. We used our maximal induction protocol, 100 μL of tamoxifen at 10 mg/mL over 5 consecutive days, for cKitCreERT2-GCaMP6f, Myh11CreERT2-GCaMP6f, and PdgfrαCreERTM -GCaMP6f mice. Due to the paucity of recombined cells in the cKitCreERT2-ROSA26mT/mG reporter mice, we used our maximal tamoxifen induction protocol for cKitCreERT2-ChR2 mice as this still resulted in the ability to excite single recombined cells. Myh11CreERT2-ChR2/tdTomato mice were induced with one 100 μL i.p. injection of tamoxifen at 0.2 mg/mL while PdgfrαCreERTM-ChR2/tdTomato were induced with 1 injection at 0.4 mg/mL tamoxifen to get mosaic induction sufficient for single cell stimulation. All mice, regardless of induction duration, were given at least 2 weeks to recover following tamoxifen injection.

Revised version: Mice harboring PdgfrαCreERTM, PdgfrβCreERT2, Myh11CreERT2, and cKitCreERT2 were crossed with ROSA26mT/mG mice to generate PdgfrαCreERTM-ROSA26mT/mG, PdgfrβCreERT2-ROSA26mT/mG, Myh11CreERT2-ROSA26mT/mG, and cKitCreERT2-ROSA26mT/mG mice, respectively. The resulting iCre-ROSA26mT/mG mice were induced with tamoxifen 2-4 weeks after weaning for confocal imaging. Mice aged 2-6 months were injected with tamoxifen for contraction studies, FACs analysis, GCaMP6f imaging, and Chr2 induction. Tamoxifen induction was performed via consecutive 100 μL i.p. injections of tamoxifen ranging from 1 to 5 days at concentrations ranging from 0.2 -10 mg/mL in safflower oil, using a titrated induction protocol to determine the extent of recombination in specific cell populations. We used our maximal induction protocol, 100 μL of tamoxifen at 10 mg/mL over 5 consecutive days, for cKitCreERT2-GCaMP6f, Myh11CreERT2-GCaMP6f, and PdgfrαCreERTM -GCaMP6f mice. Due to the paucity of recombined cells in the cKitCreERT2-ROSA26mT/mG reporter mice, we used our maximal tamoxifen induction protocol for cKitCreERT2-ChR2/tdTomato mice as this still resulted in the ability to excite single recombined cells. Myh11CreERT2-ChR2/tdTomato mice were induced with one 100 μL i.p. injection of tamoxifen at 0.2 mg/mL while PdgfrαCreERTM-ChR2/tdTomato were induced with 1 injection at 0.4 mg/mL tamoxifen to get mosaic induction sufficient for single cell stimulation. All mice, regardless of induction duration, were given at least 2 weeks to recover following tamoxifen injection.

Lymphatic Vessel Isolation

Earlier version: We utilized both popliteal and inguinal-axillary lymphatic collecting vessels (IALVs) in this study, which were isolated as described previously (Zawieja et al., 2018a). In brief, mice were anaesthetized with a cocktail of 100/10 mg/mL ketamine/xylazine and shaved along the flank or the legs for IALVs and popliteal cLVs respectively. The IALV (also referred to as the flank cLV) is located adjacent to the thoracoepigastric vein and connects the inguinal and axillary lymph nodes. A cut was made along the dorsal midline and the skin retracted and pinned out to reveal the thoracoepigastric vascular bed. The thoracoepigastric vascular bed and connected perivascular adipose containing the IALVs vessel was dissected out and pinned onto a Sylgard coated dish in Krebs buffer. Popliteal lymphatic vessels were exposed through a superficial incision in the leg, removed and transferred to the Krebs-albumin filled dissection chamber. After removal, the vessel was carefully cleaned of adipocytes and excess matrix using fine forceps and scissors through micro-dissection. For immunofluorescence, sections containing 2-3 valves were isolated, while smaller IALV sections consisting of 1-2 valves were isolated for GCaMP6f Ca2+ imaging. Similarly, popliteal cLVs were isolated (Castorena-Gonzalez et al., 2018a) following an incision along the skin overlying the saphenous vein for contractile function analysis and for ChR2 optogenetic depolarization experiments.

Revised version: We utilized both popliteal and inguinal-axillary lymphatic collecting vessels (IALVs) in this study, which were isolated as described previously (Zawieja et al., 2018a). In brief, mice were anaesthetized with a cocktail of 100/10 mg/mL ketamine/xylazine and shaved along the flank or the legs for IALVs and popliteal cLVs respectively. The IALV (also referred to as the flank cLV) is located adjacent to the thoracoepigastric vein and connects the inguinal and axillary lymph nodes. A cut was made along the dorsal midline and the skin retracted and pinned out to reveal the thoracoepigastric vascular bed. The thoracoepigastric vascular bed and connected perivascular adipose containing the IALV vessels was dissected out and pinned onto a Sylgard coated dish in Krebs buffer. Popliteal lymphatic vessels were exposed through a superficial incision in the leg, removed and transferred to the Krebs-albumin filled dissection chamber. After removal, the vessel was carefully cleaned of adipocytes and excess matrix using fine forceps and scissors through micro-dissection. For immunofluorescence, sections containing 2-3 valves were isolated, while shorter IALV sections consisting of 1-2 valves were isolated for GCaMP6f Ca2+ imaging. Similarly, popliteal cLVs were isolated (Castorena-Gonzalez et al., 2018a) following an incision along the skin overlying the saphenous, removed and transferred to the Krebs-albumin filled dissection chamber; these vessels were used for ChR2 optogenetic depolarization experiments.

Lymphatic Vessel Isobaric Function

Earlier version: PdgfrαCreERTM mice were crossed with Ano1fl/fl, Cx45fl/fl, and Cav1.2fl/fl mice to generate PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice. These mice and their respective fl/fl controls were injected with tamoxifen as described above for 5 days and given two weeks to recover. The popliteal vessels were isolated, cleaned, and prepared for isobaric contractile tests as previously reported (Davis et al., 2023). Once equilibrated, inner diameter was tracked over a physiological pressure range (stepped from 3 to 2, 1, 0.5, 3, 5, 8, and 10 cmH2O) with 2min of recording at each pressure. Following the pressure step protocol the vessels were equilibrated in with Ca2+-free Krebs buffer (3mM EGTA) and diameter at each pressure recorded under passive conditions (DMAX). The contractile parameters end diastolic diameter (EDD), end systolic diameter (ESD), and contraction frequency (FREQ) were recorded with a custom LabVIEW program and the following contractile parameters assessed:

Revised version: PdgfrαCreERTM mice were crossed with Ano1fl/fl, Cx45fl/fl, and Cav1.2fl/fl mice to generate PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice. These mice and their respective fl/fl controls were injected with tamoxifen as described above for 5 days and given two weeks to recover. The popliteal vessels were isolated, cleaned, and prepared for isobaric contractile tests as previously reported (Davis et al., 2023a). Once equilibrated, inner diameter was tracked over a physiological pressure range (stepped from 3 to 2, 1, 0.5, 3, 5, 8, and 10 cmH2O) with 2min of recording at each pressure. Following the pressure step protocol the vessels were equilibrated in with Ca2+-free Krebs buffer (3mM EGTA) and diameter at each pressure recorded under passive conditions (DMAX). The contractile parameters end diastolic diameter (EDD), end systolic diameter (ESD), and contraction frequency (FREQ) were recorded with a custom LabVIEW program and the following contractile parameters assessed:

Contraction Amplitude (AMP) = EDD−ESD

Normalized Contraction Amplitude = ((EDD−ESD)/DMAX) × 100

Ejection Fraction (EF) = (EDD2−ESD2)/EDD2

Fractional Pump Flow (FPF) = EF × FREQ

Tone = ((DMAX−EDD)/DMAX) × 100

Methylene Blue Staining

Earlier version: Isolated IALVs sections were transferred into a Krebs-BSA buffer filled 3-mL observation chamber, with a cover slip bottom, and cannulated onto two glass micropipettes (30-80 μm, outer diameter) held in place by pipette holders on a Burg-style V-track mounting system. The pipette holders were attached to a 3-way valve stop cock with polyethylene tubing filled with Krebs-BSA buffer. Vessels were pressurized to approximately 5 cmH2O by raising the 3-way valve and the vessels were stretched to remove any slack. For methylene blue staining, IALVs from wild type C57Bl6 mice were stained with 50 μM methylene blue in Krebs-BSA buffer for two hours at room temperature and covered in foil to limit light-induced phototoxicity. After the staining period, the vessel chambers were washed three times with Ca2+ free PSS to remove methylene blue. Brightfield images and manual Z-stack videos were collected on an inverted Leica DMi1 4X or 20X air objective, or a Leica DMi8 with a 25X water objective or an inverted DMi8 using a Leica Flexacam C1 color camera for image acquisition. Some Methylene blue images were also collected using a color Nikon DS-Fi3 camera. The collected z-stacks were analyzed using Image J and the “Stack Focuser” plugin (https://imagej.nih.gov/ij/plugins/stack-focuser.html). To accentuate the methylene blue stained cells, the color image stack was split into red, green, and blue channel stacks. The blue channel stack was then divided by the green channel stack using the “Image Calculator” function. The resulting 32-bit image was then converted into 16-bit image to permit the use of the Stack Focuser plugin with the ‘n kernel value’ set to 11.

Revised version: Isolated IALVs sections were transferred into a Krebs-BSA buffer filled 3-mL observation chamber, with a cover slip bottom, and cannulated onto two glass micropipettes (30-80 μm, outer diameter) held in place by pipette holders on a Burg-style V-track mounting system. The pipette holders were attached to a 3-way valve stop cock with polyethylene tubing filled with Krebs-BSA buffer. Vessels were pressurized to approximately 5 cmH2O by raising the 3-way valve and the vessels were stretched to remove any slack. For methylene blue (Sigma, M9140) staining, IALVs from wild type C57Bl6 mice were stained with 50 μM methylene blue in Krebs-BSA buffer for two hours at room temperature and covered in foil to limit light-induced phototoxicity. After the staining period, the vessel chambers were washed three times with Ca2+ free PSS to remove methylene blue. Brightfield images and manual Z-stack videos were collected on an inverted Leica DMi1 4X or 20X air objective, or a Leica DMi8 with a 25X water objective or an inverted DMi8 using a Leica Flexacam C1 color camera for image acquisition. Some methylene blue images were also collected using a color Nikon DS-Fi3 camera. The collected z-stacks were analyzed using Image J and the “Stack Focuser” plugin (https://imagej.nih.gov/ij/plugins/stack-focuser.html). To accentuate the methylene blue stained cells, the color image stack was split into red, green, and blue channel stacks. The blue channel stack was then divided by the green channel stack using the “Image Calculator” function. The resulting 32-bit image was then converted into 16-bit image to permit the use of the Stack Focuser plugin with the ‘n kernel value’ set to 11.

Fluorescence Confocal Imaging

Earlier version: IALVs vessels from each respective iCre-ROSA26mT/mG mouse were prepared in a similar manner (excluding the addition of methylene blue). We performed confocal imaging to acquire z-stacks of 7-10 overlapping regions of interests to allow for manual stitching, with 1 μM z-steps (20Χ) or 0.5 μM steps at 40X. We imaged through to the midpoint of the vessel except when imaging the valve interstitial cells, in which case the entire vessel was imaged. Max projections were made using FIJI. Following live imaging, the vessels were pressurized to 5 cmH2O and fixed with 4% paraformaldehyde for 30 min at room temperature. IALVs were then washed with PBS containing 0.1% Triton X-100 (PBST) 3 times and blocked for a minimum of 2 hr with Blockaid® (B-10710, ThermoFisher Scientific). IALVs were then stained with the corresponding primary antibodies in BlockAid® Solution: anti-smooth muscle actin (SMA) 1:500 (Sigma, A2547), anti-GFP 1:200 (ThermoFisher, A11122), anti-cKit 1:100 (Cell Signaling, 3074), anti-Vimentin 1:100 (Thermofisher, OMA1-06001), anti-desmin 1:200 (Invitrogen, PA5-16705), anti-GFP 1:200 (Abcam, ab13970, anti-CD34 1:200 (Invitrogen, 14-0341-82), anti-PDGFRΑ 1:200 (R&DSystems, AF1062), anti-PDGFRβ 1:200 (eBiosciences, 14-1402-82), anti-calponin 1:500 (Abcam, AB46794), anti-MYH11 1:500 (Abcam, AB124679), anti-Sca1 1:200 (Biolegend, 108101). IALVs were then washed in PBS and incubated overnight with the corresponding donkey secondary antibodies (ThermoFisher®) at 1:200. After a final wash, IALVs were re-cannulated and pressurized for imaging using the aforementioned spinning disk confocal and Hamamatsu Orca Flash4 camera using a 20X air objective (Olympus UplanApo, 0.75) or 40X (Olympus UApo A340, 1.15) water objective. Images were taken as described above, and the resulting stacks were turned into a max projection using FIJI. Colocalization analysis of the max projections of CD34 and PDGFRα was performed using the BIOP JACoP colocalization plugin (Bolte and Cordelieres, 2006) with both Pearson’s and Mander’s coefficients reported.

Revised version: IALVs vessels from each respective iCre-ROSA26mT/mG mouse were prepared in a similar manner (excluding the addition of methylene blue). We performed confocal imaging to acquire z-stacks of 7-10 overlapping regions of interests to allow for manual stitching, with 1 μm z-steps at (20Χ) or 0.5 μm steps at 40X. We imaged through to the midpoint of the vessel except when imaging the valve interstitial cells, in which case the entire vessel was imaged. Max projections were made using FIJI. Following live imaging, the vessels were pressurized to 5 cmH2O and fixed with 4% paraformaldehyde for 30 min at room temperature. IALVs were then washed with PBS containing 0.1% Triton X-100 (PBST) 3 times and blocked for a minimum of 2 hr with Blockaid® (B-10710, ThermoFisher Scientific). IALVs were then stained with the corresponding primary antibodies in BlockAid® Solution: anti-smooth muscle actin (SMA) 1:500 (Sigma, A2547), anti-GFP 1:200 (ThermoFisher, A11122), anti-cKIT 1:100 (Cell Signaling, 3074), anti-VIMENTIN 1:100 (Thermofisher, OMA1-06001), anti-desmin 1:200 (Invitrogen, PA5-16705), anti-GFP 1:200 (Abcam, ab13970, anti-CD34 1:200 (Invitrogen, 14-0341-82), anti-PDGFRΑ 1:200 (R&DSystems, AF1062), anti-PDGFRβ 1:200 (eBiosciences, 14-1402-82), anti-calponin 1:500 (Abcam, AB46794), anti-MYH11 1:500 (Abcam, AB124679), anti-Sca1 1:200 (Biolegend, 108101). IALVs were then washed in PBS and incubated overnight with the corresponding donkey secondary antibodies (ThermoFisher®) at 1:200. After a final wash, IALVs were re-cannulated and pressurized for imaging using the spinning disk confocal microscope and Hamamatsu Orca Flash4 camera using a 20X air objective (Olympus UplanApo, 0.75) or 40X (Olympus UApo A340, 1.15) water objective. Images were taken as described above, and the resulting stacks were turned into a max projection using FIJI. Colocalization analysis of the max projections of CD34 and PDGFRα was performed using the BIOP JACoP colocalization plugin (Bolte and Cordelieres, 2006) with both Pearson’s and Mander’s coefficients reported.

LMC Dissociation and FACS Collection-

Earlier version: IALVs vessels PdgfrαCreERTM-ROSA26mT/mG, PdgfrβCreERT2-ROSA26mT/mG, Myh11CreERT2-ROSA26mT/mG, Macgreen, and Prox1-eGFP mice were dissected and cleaned of excess adventitia and adipose tissue in Krebs buffer. Isolated vessels were then transferred into a low Ca2+ PSS solution supplemented with 0.1 mg/mL bovine serum albumin (BSA, Amersham Life Science, Arlington Heights, IL). Primary LMCs were collected by enzymatic dissociation of IALVs. The dissected vessels were cleaned in room temperature Krebs-BSA buffer and then transferred into a 1-mL tube of low-Ca2+ PSS on ice, washed, and equilibrated for 10 min. Vessels were then digested in low-Ca2+ PSS with 26 U/mL papain (Sigma, St. Louis, MO) and 1 mg/mL dithioerythritol for 30 min at 37°C with gentle agitation every few minutes. This solution was then decanted and replaced with low-Ca2+ PSS with containing 1.95 collagenase H (U/mL, Sigma), 1.8 mg/mL collagenase F (Sigma), and 1mg/mL elastase (Worthington LS00635) and incubated for 3 - 5 min at 37°C. The mixture was then spun down at 1000 rpm for 4 min, the digestion buffer removed, and replaced with low-Ca2+ PS. This process was repeated twice to remove residual digestion buffer. The vessel was then triturated with a fire-polished Pasteur pipette to dissociate the cells into a single cell suspension, passed through a Falcon cap strainer (35 μm), and resuspended in ice-cold low-Ca2+ PSS for sorting. For iCre-ROSA26mT/mG mice, GFP+RFP- cells or GFP+ cells from Macgreen and Prox1-eGFP mice were then FACS purified straight into RNA isolation buffer for RT-PCR analysis. FACs was performed with a Beckman-Coulter MoFlo XDP instrument using an excitation laser (488 m) and emission filter (530/40 m). Sorting was performed using 70-µm nozzle at a sheath pressure of 45 p.s.i. and sort rate of 100 events/s and with an efficiency of >90%. To maximize cell yield, we isolated both the left and right full-length IALVs vessels from 2 mice for digestions and subsequent FACS collection. For Myh11CreERT2-ROSA26mT/mG and PdgfrαCreERTM-ROSA26mT/mG, the yield averaged 1000-2000 cells per mouse. For Prox1-eGFP mice, LEC yield was typically 1500-2000 cells per mouse.

Revised version: IALVs vessels PdgfrαCreERTM-ROSA26mT/mG, PdgfrβCreERT2-ROSA26mT/mG, Myh11CreERT2-ROSA26mT/mG, Macgreen, and Prox1-eGFP mice were dissected and cleaned of excess adventitia and adipose tissue in Krebs buffer. Isolated vessels were then transferred into a low Ca2+ PSS solution supplemented with 0.1 mg/mL bovine serum albumin (BSA, Amersham Life Science, Arlington Heights, IL). Primary LMCs were collected by enzymatic dissociation of IALVs. The dissected vessels were cleaned in room temperature Krebs-BSA buffer and then transferred into a 1-mL tube of low-Ca2+ PSS on ice, washed, and equilibrated for 10 min. Vessels were then digested in low-Ca2+ PSS with 26 U/mL papain (Sigma, P4762) and 1 mg/mL dithioerythritol for 30 min at 37°C with gentle agitation every few minutes. This solution was then decanted and replaced with low-Ca2+ PSS with containing 1.95 collagenase H (U/mL, Sigma), 1.8 mg/mL collagenase F (Sigma), and 1mg/mL elastase (Worthington LS00635) and incubated for 3 - 5 min at 37°C. The mixture was then spun down at 1000 rpm for 4 min, the digestion buffer removed, and replaced with low-Ca2+ PS. This process was repeated twice to remove residual digestion buffer. The vessel was then triturated with a fire-polished Pasteur pipette to dissociate the cells into a single cell suspension, passed through a Falcon cap strainer (35 μm), and resuspended in ice-cold low-Ca2+ PSS for sorting. For iCre-ROSA26mT/mG mice, GFP+RFP- cells or GFP+ cells from Macgreen and Prox1-eGFP mice were then FACS purified straight into RNA isolation buffer for RT-PCR analysis. FACs was performed with a Beckman-Coulter MoFlo XDP instrument using an excitation laser (488 m) and emission filter (530/40 m). Sorting was performed using 70-µm nozzle at a sheath pressure of 45 p.s.i. and sort rate of 100 events/s and with an efficiency of >90%. To maximize cell yield, we isolated both the left and right full-length IALVs vessels from 2 mice for digestions and subsequent FACS collection. For Myh11CreERT2-ROSA26mT/mG and PdgfrαCreERTM-ROSA26mT/mG, the yield averaged 1,000-2,000 cells per mouse. For Prox1-eGFP mice, LEC yield was typically 1,500-2,000 cells per mouse.

RT-PCR Profiling of FACS Purified Cells

Total RNA was extracted from FACS purified GFP+ cells from the isolated IALVs vessels using the Arcturus PicoPure RNA isolation kit (ThermoFisher Scientific, Waltham, MA) per the listed instructions. Prior to elution in 20 μl of water, on-column DNAse digestion (Qiagen, Valencia, CA) was performed to ensure removal of genomic DNA contaminants. RNA was converted into cDNA using SuperScript III First-Strand Synthesis System (Thermo Fisher Scientific, Waltham, MA) using oligo (dT) and random hexamer priming following the manufacturer’s protocol. Each RT reaction used approximately 50-100 cells worth of RNA based on the sorted cells count number. Our PCR reaction mixture contained first-strand cDNA as the template, 2 mM MgCl2, 0.25 μM primers, 0.2 mM deoxynucleotide triphosphates; and GoTaq® Flexi DNA polymerase (Promega, Madison, WI). The PCR program comprised an initial denaturation step at 95°C for four min; followed by 35 repetitions of the following cycle: denaturation (94° C, 30 s), annealing (58° C, 30 s), and extension (72° C, 30 s). This was followed by a final elongation step for 5 min at 72° C. PCR amplification products were separated on a 2% agarose gel by electrophoresis, stained with SYBR-Safe (Thermo Fisher Scientific, Waltham, MA), and visualized by UV trans-illumination. All primers were designed to amplify an intron-spanning region. Endpoint RT-PCR Primer sequences, amplicon size, accession numbers, and source are listed in Table 1.

scRNASeq Analysis of Mouse IALVs

Earlier version: For scRNASeq analyses of isolated IALVs we used a total of 10 Rosa26mTmG mice, without Cre and without tamoxifen treatment, with equivalent representation of sex (5 males and 5 females). Full length IALVs from both the left and right side of each ROSA26mTmG mouse were isolated and cleaned of excessive matrix and adipose tissue. Isolated vessels were digested into single cell suspensions as described above and the cells were kept on ice following single cell suspension until all the tissues had been processed. Cells from all vessels were combined and sorted for tdTomato expression to remove debris and concentrate the cells for downstream single cell 3’ RNA-Seq libraries creation with 10x Genomics Chromium Chip and Chromium Next GEM Single Cell 3’ RNA-Seq reagents. Samples were sequenced with the NovaSeq 6000 S4-PE100 flow cell.

Revised version: For scRNASeq analyses of isolated IALVs we used a total of 10 Rosa26mTmG mice, without Cre expression and without tamoxifen treatment, with equivalent representation of sex (5 males and 5 females with ages between 10-12 months). Full length IALVs from both the left and right side of each ROSA26mTmG mouse were isolated and cleaned of excessive matrix and adipose tissue. Isolated vessels were digested into single cell suspensions as described above and the cells were kept on ice following single cell suspension until all the tissues had been processed. Cells from all vessels were combined and sorted for tdTomato expression to remove debris and concentrate the cells for downstream single cell 3’ RNA-Seq libraries creation with 10x Genomics Chromium Chip and Chromium Next GEM Single Cell 3’ RNA-Seq reagents. Samples were sequenced with the NovaSeq 6000 S4-PE100 flow cell.

Earlier version: Mus musculus genome GRCm39 and annotation GTF (v106) from Ensembl (https://useast.ensembl.org/Mus_musculus/Info/Index) were used to build the reference index and the reads were processed using Cell Ranger (v7.0.1; Zheng et al., 2017) with default parameters. The quality control and filtering steps were performed using R (v4.2.1; https://www.r-project.org/). Ambient RNA was removed from the Cell Ranger output with SoupX (Young and Behjati, 2020). Doublet score for each cell was estimated using scDBlFinder (v1.12.0; Germain et al., 2021). Non-expressed genes (sum zero across all samples) and low-quality cells (>10% mitochondrial genes, < 500 genes, < 1,000 UMIs per cell and doublet score <0.5) were removed with custom R scripts. Cells passing filtering were normalized/scaled (SCTransformation), dimensionally reduced (t-distributed stochastic neighbor embedding (t-SNE) and uniform manifold approximation and projection (UMAP)) clustered, and hierarchically analyzed with Seurat (Hao et al., 2021; Hao et al., 2024) with default parameters. Marker gene expression profile on cell clusters and gene co-expression was visualized using Seurat and ShinyCell R application (Ouyang et al., 2021). The full scRNAseq raw dataset has been uploaded to the NIH GEO under the accession number GSE277843.. Differential gene expression within subclusters of LECs, LMCs, and AdvCs was performed using Seurat’s “Find Markers” function and with a minimum of either 40% or 50% cell expression and log fold change minimum of 1. In the volcano plot for LEC subcluster 8 differential gene expression, listed genes were had a cutoff of a log fold change of 2 or −2 to be displayed on the plot.

Revised version: Mus musculus genome GRCm39 and annotation GTF (v106) from Ensembl (https://useast.ensembl.org/Mus_musculus/Info/Index) were used to build the reference index and the reads were processed using Cell Ranger (v7.0.1; Zheng et al., 2017) with default parameters. The quality control and filtering steps were performed using R (v4.2.1; https://www.r-project.org/). Ambient RNA was removed from the Cell Ranger output with SoupX (Young and Behjati, 2020). Doublet score for each cell was estimated using scDBlFinder (v1.12.0; Germain et al., 2021). Non-expressed genes (sum zero across all samples) and low-quality cells (>10% mitochondrial genes, < 500 genes, < 1,000 UMIs per cell and doublet score <0.5) were removed with custom R scripts. Of 10,188 cells, 7,435 passed our inclusion criteria and included three dominant clusters including LECs (2,962 cells), LMCs (978 cells), and fibroblast (2,261 cells) with the remaining cells comprising immune cells (1,147 cells) and some mammary epithelial cell contamination (87 cells). Cells passing filtering were normalized/scaled (SCTransformation), dimensionally reduced (t-distributed stochastic neighbor embedding (t-SNE) and uniform manifold approximation and projection (UMAP)) clustered, and hierarchically analyzed with Seurat (Hao et al., 2021; Hao et al., 2024) with default parameters. To select the optimal cluster resolution, we used Clustree with various resolutions. We examined the resulting tree to identify a resolution where the clusters were well-separated and biologically meaningful, ensuring minimal merging or splitting at higher resolutions. Our goal was to find a resolution that captured relevant cell subpopulations while maintaining distinct clusters without excessive fragmentation. Initial clustering of the entire population of cells was done at a resolution of 0.8 and 18 PCs to achieve the UMAP of 19 cell clusters as shown in Figure 5a. We used a resolution of 0.5 for sub-clustering LMCs (original groups 5 and 6), 0.87 for LECs (original groups 0,1,2,11), and 1.0 for fibroblasts (3,7,8,9,10,13). Marker gene expression profile on cell clusters and gene co-expression was visualized using Seurat and ShinyCell R application (Ouyang et al., 2021). The full scRNAseq raw dataset has been uploaded to the NIH GEO under the accession number GSE277843. Differential gene expression within subclusters of LECs, LMCs, and AdvCs was performed using Seurat’s “Find Markers” function and with a minimum of either 40% or 50% cell expression and average log fold change (Log2FC) minimum of 0.5 or 1. When assessing the LMC IP3 receptor genes Itpr1-3 and Itprid2 a percent cell expression of 40% and Log2FC of 0.25 was used. In the volcano plot for LEC subcluster 8 differential gene expression, listed genes had a cutoff of a Log2FC of 2 or −2 to be displayed on the plot.

Ex vivo Ca2+ imaging with the genetically encoded GCaMP6f Indicator

Earlier version: cKitCreERT2, Myh11CreERT2, and PdgfrαCreERTM mice were crossed with GCaMP6f mice in a similar manner as described for ROSA26mT/mG. cKitCreERT2-GCaMP6f, PdgfrαCreERTM-GCaMP6f, and Myh11CreERT2-GCaMP6f were induced with tamoxifen (10 mg/mL) for 5 consecutive days by i.p. injection. IALVs isolated from cKitCreERT2-GCaMP6f, PdgfrαCreERTM-GCaMP6f, and Myh11CreERT2-GCaMP6f were cannulated as described above. The cannulated vessel, with micropipette holders, observation chamber and V-track mounting system, was transferred to the stage of the spinning disk confocal with a Prime95B scMOS camera (Photometrics), a Cascade II EMCCD (Photometrics), or an Ixon888 EMCCD camera (Andor) for Ca2+ imaging (Castorena-Gonzalez et al., 2018b). Pressures for the input and output cannula were connected to a T-junction which was set briefly to 8 cmH2O and the vessel lengthened to remove axial slack. A peristaltic pump maintained constant perfusion of the observation chamber with Krebs buffer at a rate of 0.5 mL/min while the vessel equilibrated at 37°C for 30-60 min with pressures set to 3 cmH2O. Spontaneous contractions were allowed to stabilize over a period of 30 min and then were blunted with 2 μM wortmannin to limit movement associated with contractions during Ca2+ imaging. A Windows-based computer was used to digitize the pressure transducer signals and video image of the vessel from a firewire camera at 30-40 Hz (Davis et al., 2012). A custom-written LabVIEW program (National Instruments; Austin, TX) detected the inner diameter of the vessel from the video (Davis et al., 2011). Once contractions were <5 µm in amplitude, Ca2+ recordings were made at 20FPS for 20-40 s.

Revised version: cKitCreERT2, Myh11CreERT2, and PdgfrαCreERTM mice were crossed with GCaMP6f mice in a similar manner as described for ROSA26mT/mG. cKitCreERT2-GCaMP6f, PdgfrαCreERTM-GCaMP6f, and Myh11CreERT2-GCaMP6f were induced with tamoxifen (10 mg/mL) for 5 consecutive days by i.p. injection. IALVs isolated from cKitCreERT2-GCaMP6f, PdgfrαCreERTM-GCaMP6f, and Myh11CreERT2- GCaMP6f were cannulated as described above. The cannulated vessel, with micropipette holders, observation chamber and V-track mounting system, was transferred to the stage of the spinning disk confocal with a Prime95B scMOS camera (Photometrics), a Cascade II EMCCD (Photometrics), or an Ixon888 EMCCD camera (Andor) for Ca2+ imaging (Castorena-Gonzalez et al., 2018b). Pressures for the input and output cannula were connected to a T-junction which was set briefly to 8 cmH2O and the vessel lengthened to remove axial slack. A peristaltic pump maintained constant perfusion of the observation chamber with Krebs buffer at a rate of 0.5 mL/min while the vessel equilibrated at 37°C for 30-60 min with pressures set to 3 cmH2O. Spontaneous contractions were allowed to stabilize over a period of 30 min to verify normal function and then were blunted with 2 μM wortmannin to limit movement associated with contractions during Ca2+ imaging. A Windows-based computer was used to digitize the pressure transducer signals and video image of the vessel from a firewire camera at 30-40 Hz (Davis et al., 2012). A custom- written LabVIEW program (National Instruments; Austin, TX) detected the inner diameter of the vessel from the video (Davis, 2005). Once contractions were <5 µm in amplitude, Ca2+ recordings were made at 20FPS for 20-40 s.

Ca2+ Imaging and Analysis in IALVs Over the Contraction Cycle

Background noise was determined by using the histogram feature of FIJI in a rectangle in a region of the field of view without sample. This value was subtracted from the entire field of view. In some cases, the vessel movement due to contraction was offset with video stabilization with the FIJI plugin Image Stabilizer. A max projection was used to create non-overlapping ROIs of GCaMP6f+ cells for each iCre- GCaMp6f IALV. From these cell ROIs, the “reslice z” function was used to create pseudo-linescan STMs, which were divided by their baseline values to obtain F/F0 values for each individual cell. At least 3 cells, except in the case of 1 cKitCreERT2-GCaMp6f IALV, in which only two cells were observed, were analyzed in this manner for each vessel segment. Max projections of the image stack were then used to create non-overlapping cell masks of 3-5 muscle cells per field of view of one vessel. Ca2+ traces for those cells contained 5-10 contraction cycles and Ca2+ transients and were characterized for peak intensity (expressed as a baseline-referenced ratio, F/F0), frequency, and duration in seconds.

Analysis of Subcellular Ca2+ Transients in Myh11CreERT2-GCaMP6f IALVs

Earlier version: For Myh11CreERT2-We performed Ca2+ imaging as above in the presence of 1 µM nifedipine to stop the “Ca2+ flashes” associated with APs (Zawieja et al., 2018a) and focus on the subcellular activity at 3 different experimental pressures of 0.5, 2, and 5 cmH2O. For this protocol, we used a particle analysis approach to analyze all Ca2+ transients in the field of view. Ca2+ transients in intact vessels were quantified by particle analysis as previously described (Drumm et al., 2017; Drumm et al., 2019). Movies of Ca2+ transients in intact vessels were imported into custom built Volumetry software (version G8d) and background subtracted. Movies were smoothed using a Gaussian filter: 1.5 x 1.5 mM, StdDev 1.0). Raw Ca2+ transients were converted to Ca2+ particles (PTCLs) using a flood-fill algorithm as previously described (Drumm et al., 2017; Drumm et al., 2019). PTCLs <10 µM2 were rejected to facilitate the removal of noise and then the total PTCL area and PTCL count could be tabulated for each recording.

Revised version: For Myh11CreERT2-We performed Ca2+ imaging as above in the presence of 1 µM nifedipine to stop the “Ca2+ flashes” associated with APs (Zawieja et al., 2018a) and focus on the subcellular activity at 3 different experimental pressures of 0.5, 2, and 5 cmH2O. For this protocol, we used a particle analysis approach to analyze all Ca2+ transients in the field of view. Ca2+ transients in intact vessels were quantified by particle analysis as previously described (Drumm et al., 2017; Drumm et al., 2019b). Movies of Ca2+ transients in intact vessels were imported into custom built Volumetry software (version G8d) and background subtracted. Movies were smoothed using a Gaussian filter: 1.5 x 1.5 mM, StdDev 1.0). Raw Ca2+ transients were converted to Ca2+ particles (PTCLs) using a flood-fill algorithm as previously described (Drumm et al., 2017; Drumm et al., 2019b). PTCLs <10 µM2 were rejected to facilitate the removal of noise and then the total PTCL area and PTCL count could be tabulated for each recording.

Membrane Potential Recordings in IALVs

Earlier version: Mouse IALVs were isolated and cleaned as described above. IALVs were pressurized in our isobaric myography apparatus and allowed to equilibrate to ensure typical contractile activity was evident. A bolus of wortmannin at 2 µM was then applied to the bath to blunt contraction amplitude below 5 microns. Intracellular recordings of lymphatic muscle were made with microelectrodes (250–300 MΩ) filled with 1 M KCl and an SEC-05x amplifier (NPI) connected to a Grass S48 stimulator, viewed with a Tektronix TDS3052 digital oscilloscope. Membrane potential and diameter were simultaneously recorded using a custom Labview program. Membrane potential and APS were allowed to stabilize and then pressure was slowly raised from 0.5 cmH2O to 2 cmH2O and then 5cmH2O. In some cases, the electrode dislodged due to the intrinsic contractions of the vessel or wall displacement as pressure was modulated. In these situations, we attempted to re-impale the cell or one of the neighboring cells. Only vessels in which a recording with a minimum of 3 stable APs was successfully made at 2 of the 3 experimental pressures were used for subsequent analysis.

Revised version: Mouse IALVs were isolated and cleaned as described above. IALVs were pressurized in our isobaric myography apparatus and allowed to equilibrate to ensure typical contractile activity was evident. A bolus of wortmannin at 2 µM was then applied to the bath to blunt contraction amplitude below 5 µm. Intracellular recordings of lymphatic muscle were made with microelectrodes (250–300 MΩ) filled with 1 M KCl and an SEC-05x amplifier (NPI) connected to a Grass S48 stimulator, viewed with a Tektronix TDS3052 digital oscilloscope. Membrane potential and diameter were simultaneously recorded using a custom Labview program. Membrane potential and APS were allowed to stabilize and then pressure was slowly raised from 0.5cmH2O to 2cmH2O and then 5cmH2O. In some cases, the electrode dislodged due to the intrinsic contractions of the vessel or wall displacement as pressure was modulated. In these situations, we attempted to re-impale the cell or one of the neighboring cells. Only vessels in which a recording with a minimum of 3 stable APs was successfully recorded at 2 of the 3 experimental pressures were used for subsequent analysis.

We also confirmed LMC impalement using microelectrode filled with 1 M KCl and (100 µg/ml) Biocytin-AF488 (A12924, ThermoFisher) to label impaled cells that displayed APs, over a 10-minute recording period. Following the impalement and loading with Biocytin-AF488 the vessel was transferred to our imaging apparatus for confocal imaging and 3D reconstruction using the Andor Dragonfly 200 and IMARIS. Image stacks were taken with a 25x water objective at 0.5-micron intervals throughout the diameter of the vessel.

Light Activation of ChR2 to stimulate Popliteal Collecting Lymphatic Vessel Contractions

Earlier version: As the IALV has a nearly continuous contractile cycle, we utilized the popliteal vessel for its much slower contraction frequency in the experiments testing our ability to trigger a propagated contraction upon stimulation of the enforced expression of ChR2. Popliteal vessels were isolated from cKitCreERT2-ChR2/tdTomato, PdgfrαCreERTM -ChR2/tdTomato, or Myh11CreERT2-ChR2/tdTomato mice as previously described (Scallan and Davis, 2013), although we intentionally retained some connective tissue and adipose tissue to ensure we had a sufficient population of recombined cells to test in the adventitia layer of the vessel. Contractions were allowed to stabilize over a 30-min equilibration period with pressure set to 3 cmH2O. If basal contraction frequency was too high, we applied pinacidil to the bath in 100 nM increments, without exceeding 600 nM, to further slow contraction frequency to around 6 contractions per minute. Pinacidil at sub 1 µM doses can slow contraction frequency without causing overt hyperpolarization of membrane potential (Davis et al., 2020). Supplemental 100 nM doses of pinacidil were applied throughout the experiment to maintain a spontaneous contraction frequency below 6 per minute to allow ample diastolic time for ChR2 stimulation. Throughout this protocol the popliteal was allowed to contract spontaneously to ensure we had not overly inhibited APs by the pacemaking cells with pinacidil. Occasionally, spontaneous contractions occurred just prior to light-evoked contractions, resulting in a potential false positive, so we performed multiple stimulations over a period of 5 - 10 min, typically waiting at least 3 s after any spontaneous contraction before stimulating. Care was made to align the light fiber in such a way that only part of the vessel would be directly illuminated and so target cells of interest would be directly activated by 473 nm light using a Laser diode (Doric LD Fiber Light Source, Quebec, Canada), through an optical probe with a 10-µm tip (Doric, OPT_200_0.22.010). To further limit the excitation field, the optical probe was coated with black acrylic paint using an eyelash brush so that the uncoated opening was ∼2-3 µm. With the probe positioned within 5 µm of one side of the vessel wall, the spread of light covered an area ∼10-100 µm wide on the back side of the vessel (depending on the diode amplitude setting). Light pulses, 200 ms in length, were triggered by a Grass S9 stimulator (Harvard Apparatus, Holliston, MA) connected to the external TTL input of the laser diode. Pulse amplitude was adjusted between 40-90 mA using the Laser Diode Module Driver (Doric). A contraction was considered to be triggered if it occurred within 50ms of stimulation. We performed photo-stimulation from 2-4 sites within each vessel, with 6-14 stimulations per site. If a photo-stimulation was triggered incidentally after the initiation of a “spontaneous contraction” it was discarded from the analysis. For Myh11CreERT2-ChR2-tdTomato 6 vessels from 3 separate mice were tested. For PdgfrαCreERTM-ChR2-tdTomato 6 vessels from 4 separate mice were tested with a max of two vessels per mouse. For cKitCreERT2-ChR2-tdTomato 7 vessels from 3 separate mice were assessed. Diameter was recorded to align photo-activation with the contraction cycle in a custom Labview program.

Revised version: As the IALV has a nearly continuous contractile cycle, we utilized the popliteal vessel for its much slower contraction frequency in the experiments testing our ability to trigger a propagated contraction upon stimulation of the enforced expression of ChR2. Popliteal vessels were isolated from cKitCreERT2-ChR2/tdTomato, PdgfrαCreERTM -ChR2/tdTomato, or Myh11CreERT2-ChR2/tdTomato mice (3-9 months of age) as previously described (Scallan and Davis, 2013), although we intentionally retained some connective tissue and adipose tissue to ensure we had a sufficient population of recombined cells to test in the adventitia layer of the vessel. Contractions were allowed to stabilize over a 30-min equilibration period with pressure set to 3 cmH2O. If basal contraction frequency was too high, we applied pinacidil to the bath in 100 nM increments, without exceeding 600 nM, to further slow contraction frequency to around 6 contractions per minute. Pinacidil at sub 1 µM doses can slow contraction frequency without causing overt hyperpolarization of membrane potential (Davis et al., 2020). Supplemental 100 nM doses of pinacidil were applied throughout the experiment to maintain a spontaneous contraction frequency below 6 per minute to allow ample diastolic time for ChR2 stimulation. Throughout this protocol the popliteal was allowed to contract spontaneously to ensure we had not overly inhibited APs by the pacemaking cells with pinacidil. Occasionally, spontaneous contractions occurred just prior to light-evoked contractions, resulting in a potential false positive, so we performed multiple stimulations over a period of 5 - 10 min, typically waiting at least 3 s after any spontaneous contraction before stimulating. Care was made to align the light fiber in such a way that only part of the vessel would be directly illuminated and so target cells of interest would be directly activated by 473 nm light using a Laser diode (Doric LD Fiber Light Source, Quebec, Canada), through an optical probe with a 10-µm tip (Doric, OPT_200_0.22.010). To further limit the excitation field, the optical probe was coated with black acrylic paint using an eyelash brush so that the uncoated opening was ∼2-3 µm. With the probe positioned within 5 µm of one side of the vessel wall, the spread of light covered an area ∼10-100 µm wide on the back side of the vessel (depending on the diode amplitude setting). Light pulses, 200 ms in length, were triggered by a Grass S9 stimulator (Harvard Apparatus, Holliston, MA) connected to the external TTL input of the laser diode. Pulse amplitude was adjusted between 40-90 mA using the Laser Diode Module Driver (Doric). A contraction was considered to be triggered if it occurred within 50ms of stimulation. We performed photo-stimulation from 2-4 sites within each vessel, with 6-14 stimulations per site. If a photo-stimulation was triggered incidentally after the initiation of a “spontaneous contraction” it was discarded from the analysis. For Myh11CreERT2-ChR2-tdTomato 6 vessels from 3 separate mice were tested. For PdgfrαCreERTM-ChR2-tdTomato 6 vessels from 4 separate mice were tested with a max of two vessels per mouse. For cKitCreERT2-ChR2-tdTomato 7 vessels from 3 separate mice were assessed. Diameter was recorded to align photo-activation with the contraction cycle in a custom Labview program.

Solutions and Chemicals

Earlier version: Krebs buffer was composed of (in mM): 146.9 NaCl, 4.7 KCl, 2 CaCl2, 1.2 MgSO4, 1.2 NaH2PO4•H2O, 3 NaHCO3, 1.5 NaHEPES, and 5 d-glucose (pH = 7.4 at 37°C). Krebs-BSA buffer was prepared with the addition of 0.5% (w/v) bovine serum albumin (BSA) while Krebs Ca2+-free replaced CaCl2 with 3mM EGTA. Tamoxifen was dissolved to 10mg/ml in a Safflower Oil-Ethanol (95%-5% v/v) solution with rocking agitation, separated into aliquots, and stored at -20 °C. Wortmannin was dissolved in DMSO to a stock solution of 1 mM. Pinacidil was dissolved in DMSO to a stock concentration of 1 µM. Nifedipine was dissolved in DMSO to a stock concentration of 1 mM. All chemicals were obtained from Sigma (St. Louis, MO), except for BSA (US Biochemicals; Cleveland, OH), MgSO4 and NaHEPES (Fisher Scientific; Pittsburgh, PA).

Revised version: Krebs buffer was composed of (in mM): 146.9 NaCl, 4.7 KCl, 2 CaCl2, 1.2 MgSO4, 1.2 NaH2PO4•H2O, 3 NaHCO3, 1.5 NaHEPES, and 5 d-glucose (pH = 7.4 at 37°C). Krebs-BSA buffer was prepared with the addition of 0.5% (w/v) bovine serum albumin (BSA) while Krebs Ca2+-free replaced CaCl2 with 3mM EGTA. Tamoxifen was dissolved to 10mg/ml in a Safflower Oil-Ethanol (95%-5% v/v) solution with rocking agitation, separated into aliquots, and stored at −20 °C. Wortmannin was dissolved in DMSO to a stock solution of 1mM. Pinacidil was dissolved in DMSO to a stock concentration of 1 µM. Nifedipine was dissolved in DMSO to a stock concentration of 1 mM. All chemicals were obtained from Sigma (St. Louis, MO), except for BSA (US Biochemicals; Cleveland, OH), MgSO4 and NaHEPES (Fisher Scientific; Pittsburgh, PA).

Statistical Tests

Earlier version: Statistical differences in the isobaric contractile tests for popliteal cLVs isolated from PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice over the various contractile parameters were assessed via 1) repeated measures two-way ANOVAs with Sidak’s multiple comparison tests performed using Prism9 (Graphpad). Data are plotted as mean ± SEM and significance determined at p < 0.05 and 0.10 > p > 0.05 were reported. Data from cLVs in which a negative tone value was recorded at any pressure, which typically indicated incomplete passivation or bubbles in the cannula, were not included in the tone analysis. Vessels that failed to contract at a given pressure had no value recorded for ejection fraction or normalized amplitude and REML mixed effects model was used in place of repeated measures 2-way ANOVA. We used a categorical Chi-squared statistical test for the experiments assessing our ability to trigger a contraction with activation of ChR2+ cells. Ca2+ particle area and frequency was compared using 1-way ANOVA with Tukey’s post-hoc test. Significance was determined at a p value of < 0.05. A mixed effects analysis with Tukey’s multiple comparison post-hoc test was used to compare AP parameters across pressure using Prism9 (Graphpad).

Revised version: Statistical differences in the isobaric contractile tests for popliteal cLVs isolated from PdgfrαCreERTM-Ano1fl/fl, PdgfrαCreERTM-Cx45fl/fl, and PdgfrαCreERTM-Cav1.2fl/fl mice over the various contractile parameters were assessed via 1) repeated measures two-way ANOVAs with Sidak’s multiple comparison tests performed using Prism9 (Graphpad). Data are plotted as mean ± SEM and significance determined at p < 0.05 and 0.10 > p > 0.05 were reported. Cre negative mice were used for controls and experiment order was not randomized aside from random mouse selection from cages housing both Cre+ or Cre- mice. Experimental sample size was determined by the results from our previous experiments assessing Cx45, Ano1, Cacna1c with Myh11CreERT2 mice. Data from cLVs in which a negative tone value was recorded at any pressure, which typically indicated incomplete relaxation or occluding bubbles in the cannula, were not included in the tone analysis. Vessels that failed to contract at a given pressure had no value recorded for ejection fraction or normalized amplitude and REML mixed effects model was used in place of repeated measures 2-way ANOVA. We used a categorical Chi-squared statistical test for the experiments assessing our ability to trigger a contraction with activation of ChR2+ cells. Ca2+ particle area and frequency were compared using 1-way ANOVA with Tukey’s post-hoc test. Significance was determined at p < 0.05. A mixed effects analysis with Tukey’s multiple comparison post-hoc test was used to compare AP parameters across pressure using Prism9 (Graphpad).